Srs2 binding to PCNA and its sumoylation contribute to RPA antagonism during the DNA damage response

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public Review):

Overall, the data presented in this manuscript is of good quality. Understanding how cells control RPA loading on ssDNA is crucial to understanding DNA damage responses and genome maintenance mechanisms. The authors used genetic approaches to show that disrupting PCNA binding and SUMOylation of Srs2 can rescue the CPT sensitivity of rfa1 mutants with reduced affinity for ssDNA. In addition, the authors find that SUMOylation of Srs2 depends on binding to PCNA and the presence of Mec1.

Comments on revisions:

I am satisfied with the revisions made by the authors, which helped clarify some points that were confusing in the initial submission.

Thank you.

Reviewer #2 (Public Review):

This revised manuscript mostly addresses previous concerns by doubling down on the model without providing additional direct evidence of interactions between Srs2 and PCNA, and that "precise sites of Srs2 actions in the genome remain to be determined." One additional Srs2 allele has been examined, showing some effect in combination with rfa1-zm2. Many of the conclusions are based on reasonable assumptions about the consequences of various mutations, but direct evidence of changes in Srs2 association with PNCA or other interactors is still missing. There is an assumption that a deletion of a Rad51-interacting domain or a PCNA-interacting domain have no pleiotropic effects, which may not be the case. How SLX4 might interact with Srs2 is unclear to me, again assuming that the SLX4 defect is "surgical" - removing only one of its many interactions.

Previous studies have already provided direct evidence for the interaction between Srs2 and PCNA through the Srs2’s PIM region (Armstrong et al, 2012; Papouli et al, 2005); we have added these citations in the text. Similarly. Srs2 associations with SUMO and Rad51 have also been demonstrated (Colavito et al, 2009; Kolesar et al, 2016; Kolesar et al., 2012), and these studies were cited in the text.

We did not state that a deletion of a Rad51-interacting domain or a PCNA-interacting domain have no pleiotropic effects. We only assessed whether these previously characterized mutant alleles could mimic srs2∆ in rescuing rfa1-zm2 defects.

We assessed the genetic interaction between slx4-RIM and srs2-∆PIM mutants, and not the physical interaction between the two proteins. As we described in the text, our rationale for this genetic test is based on that the reports that both slx4 and srs2 mutants impair recovery from the Mec1 induced checkpoint, thus they may affect parallel pathways of checkpoint dampening.

One point of concern is the use of t-tests without some sort of correction for multiple comparisons - in several figures. I'm quite sceptical about some of the p < 0.05 calls surviving a Bonferroni correction. Also in 4B, which comparison is **? Also, admittedly by eye, the changes in "active" Rad53 seem much greater than 5x. (also in Fig. 3, normalizing to a non-WT sample seems odd).

Claims made in this work were based only on pairwise comparison not multi-comparison. We have now made this point clearer in the graphs and in Method. As the values were compared between a wild-type strain and a specific mutant strain, or between two mutants, we believe that t-test is suitable for statistical analysis.

Figure 4B, ** indicates that the WT value is significantly different from that of the slx4-RIM srs2-∆PIM double mutant and from that of srs2-∆PIM single mutant. We have modified the graph to indicate the pair-wide comparison. The 5-fold change of active Rad53 levels was derived by comparing the values between the srs2∆ PIM slx4RIM-TAP double mutant and wild-type Slx4-TAP. In Figure 3, normalization to the lowest value affords better visualization. This is rather a stylish issue; we would like to maintain it as the other reviewers had no issues.

What is the WT doubling time for this strain? From the FACS it seems as if in 2 h the cells have completed more than 1 complete cell cycle. Also in 5D. Seems fast...

Wild-type W303 strain has less than 90 min doubling time as shown by many labs, and our data are consistent with this. The FACS profiles for wild-type cells shown in Figures 3C, 4C, and 5C are consistent with each other, showing that after G1 cells entered the cell cycle, they were in G2 phase at the 1-hour time points, and then a percentage of the cells exited the first cell cycle by two hours.

I have one over-arching confusion. Srs2 was shown initially to remove Rad51 from ssDNA and the suppression of some of srs2's defects by deleting rad51 made a nice, compact story, though exactly how srs2's "suppression of rad6" fit in isn't so clear (since Rad6 ties into Rad18 and into PCNA ubiquitylation and into PCNA SUMOylation). Now Srs2 is invoked to remove RPA. It seems to me that any model needs to explain how Srs2 can be doing both. I assume that if RPA and Rad51 are both removed from the same ssDNA, the ssDNA will be "trashed" as suggested by Symington's RPA depletion experiments. So building a model that accounts for selective Srs2 action at only some ssDNA regions might be enhanced by also explaining how Rad51 fits into this scheme.

While the anti-recombinase function of Srs2 was better studied, its “anti-RPA” role in checkpoint dampening was recently described by us (Dhingra et al, 2021) following the initial report by the Haber group some time ago (Vaze et al, 2002). A better understanding of this new role is required before we can generate a comprehensive picture of how Srs2 integrates the two functions (and possibly other functions). Our current work addresses this issue by providing a more detailed understanding of this new role of Srs2.

Single molecular data showed that Srs2 strips both RPA and Rad51 from ssDNA, but this effect is highly dynamic (i.e. RPA and Rad51 can rebind ssDNA after being displaced) (De Tullio et al, 2017). As such, generation of “deserted” ssDNA regions lacking RPA and Rad51 in cells can be an unlikely event. Rather, Srs2 can foster RPA and Rad51 dynamics on ssDNA. Additional studies will be needed to generate a model that integrates the anti-recombinase and the anti-RPA roles of Srs2.

As a previous reviewer has pointed out, CPT creates multiple forms of damage. Foiani showed that 4NQO would activate the Mec1/Rad53 checkpoint in G1- arrested cells, presumably because there would be singlestrand gaps but no DSBs. Whether this would be a way to look specifically at one type of damage is worth considering; but UV might be a simpler way to look. As also noted, the effects on the checkpoint and on viability are quite modest. Because it isn't clear (at least to me) why rfa1 mutants are so sensitive to CPT, it's hard for me to understand how srs2-zm2 has a modest suppressive effect: is it by changing the checkpoint response or facilitating repair or both? Or how srs2-3KR or srs2-dPIM differ from rfa1-zm2 in this respect. The authors seem to lump all these small suppressions under the rubric of "proper levels of RPA-ssDNA" but there are no assays that directly get at this. This is the biggest limitation.

CPT treatment is an ideal condition to examine how cells dampen the DNA damage checkpoint, because while most genotoxic conditions (e.g. 4NQO, MMS) induce both the DNA replication checkpoint and the DNA damage checkpoint, CPT was shown to only induced the latter (Menin et al, 2018; Minca & Kowalski, 2011; Redon et al, 2003; Tercero et al, 2003). Future studies examining 4NQO and UV conditions can further expand our understanding of checkpoint dampening in different conditions.

We have previously provided evidence to support the conclusion that srs2 suppression of rfa1-zm is partly mediated by changing checkpoint levels (Dhingra et al., 2021). We cannot exclude the possibility that the suppression may also be related to changes of DNA repair; we have now added this note in the text.

Regarding direct testing RPA levels on DNA, we have previously shown that srs2∆ increased the levels of chromatin associated Rfa1 and this is suppressed by rfa1-zm2 (Dhingra et al., 2021). We have now included chromatin fractionation data to show that srs2-∆PIM also led to an increase of Rfa1 on chromatin, and this was suppressed by rfa1-zm2 (new Fig. S2).

Srs2 has also been implicated as a helicase in dissolving "toxic joint molecules" (Elango et al. 2017). Whether this activity is changed by any of the mutants (or by mutations in Rfa1) is unclear. In their paper, Elango writes: "Rare survivors in the absence of Srs2 rely on structure-specific endonucleases, Mus81 and Yen1, that resolve toxic joint-molecules" Given the involvement of SLX4, perhaps the authors should examine the roles of structure-specific nucleases in CPT survival?

Srs2 has several roles, and its role in RPA antagonism can be genetically separated from its role in Rad51 regulation as we have shown in our previous work (Dhingra et al., 2021) and this notion is further supported by evidence presented in the current work. Srs2’s role in dissolving "toxic joint molecules” was mainly observed during BIR (Elango et al, 2017). Whether it is related to checkpoint dampening will be interesting to address in the future but is beyond of the scope of the current work that seeks to answer the question how Srs2 regulates RPA during checkpoint dampening. Similarly, determining the roles of Mus81 and Yen1 and other structural nucleases in CPT survival is a worthwhile task but it is a research topic well separated from the focus of this work.

Experiments that might clarify some of these ambiguities are proposed to be done in the future. For now, we have a number of very interesting interactions that may be understood in terms of a model that supposes discriminating among gaps and ssDNA extensions by the presence of PCNA, perhaps modified by SUMO. As noted above, it would be useful to think about the relation to Rad6.

Several studies have shown that Srs2’s functional interaction with Rad6 is based on Srs2-mediated recombination regulation (reviewed by (Niu & Klein, 2017). Given that recombinational regulation by Srs2 is genetically separable from the Srs2 and RPA antagonism (Dhingra et al., 2021), we do not see a strong rationale to examine Rad6 in this work, which addresses how Srs2 regulates RPA. With this said, this study has provided basis for future studies of possible cross-talks among different Srs2-mediated pathways.

Reviewer #3 (Public Review):

The superfamily I 3'-5' DNA helicase Srs2 is well known for its role as an anti-recombinase, stripping Rad51 from ssDNA, as well as an anti-crossover factor, dissociating extended D-loops and favoring non-crossover outcome during recombination. In addition, Srs2 plays a key role in in ribonucleotide excision repair. Besides DNA repair defects, srs2 mutants also show a reduced recovery after DNA damage that is related to its role in downregulating the DNA damage signaling or checkpoint response. Recent work from the Zhao laboratory (PMID: 33602817) identified a role of Srs2 in downregulating the DNA damage signaling response by removing RPA from ssDNA. This manuscript reports further mechanistic insights into the signaling downregulation function of Srs2.

Using the genetic interaction with mutations in RPA1, mainly rfa1-zm2, the authors test a panel of mutations in Srs2 that affect CDK sites (srs2-7AV), potential Mec1 sites (srs2-2SA), known sumoylation sites (srs2-3KR), Rad51 binding (delta 875-902), PCNA interaction (delta 1159-1163), and SUMO interaction (srs2SIMmut). All mutants were generated by genomic replacement and the expression level of the mutant proteins was found to be unchanged. This alleviates some concern about the use of deletion mutants compared to point mutations. Double mutant analysis identified that PCNA interaction and SUMO sites were required for the Srs2 checkpoint dampening function, at least in the context of the rfa1-zm2 mutant. There was no effect of this mutants in a RFA1 wild type background. This latter result is likely explained by the activity of the parallel pathway of checkpoint dampening mediated by Slx4, and genetic data with an Slx4 point mutation affecting Rtt107 interaction and checkpoint downregulation support this notion. Further analysis of Srs2 sumoylation showed that Srs2 sumoylation depended on PCNA interaction, suggesting sequential events of Srs2 recruitment by PCNA and subsequent sumoylation. Kinetic analysis showed that sumoylation peaks after maximal Mec1 induction by DNA damage (using the Top1 poison camptothecin (CPT)) and depended on Mec1. This data are consistent with a model that Mec1 hyperactivation is ultimately leading to signaling downregulation by Srs2 through Srs2 sumoylation. Mec1-S1964 phosphorylation, a marker for Mec1 hyperactivation and a site found to be needed for checkpoint downregulation after DSB induction, did not appear to be involved in checkpoint downregulation after CPT damage. The data are in support of the model that Mec1 hyperactivation when targeted to RPA-covered ssDNA by its Ddc2 (human ATRIP) targeting factor, favors Srs2 sumoylation after Srs2 recruitment to PCNA to disrupt the RPA-Ddc2-Mec1 signaling complex. Presumably, this allows gap filling and disappearance of long-lived ssDNA as the initiator of checkpoint signaling, although the study does not extend to this step.

Strengths:

(1) The manuscript focuses on the novel function of Srs2 to downregulate the DNA damage signaling response and provide new mechanistic insights.

(2) The conclusions that PCNA interaction and ensuing Srs2-sumoylation are involved in checkpoint downregulation are well supported by the data.

Weaknesses:

(1) Additional mutants of interest could have been tested, such as the recently reported Pin mutant, srs2-Y775A (PMID: 38065943), and the Rad51 interaction point mutant, srs2-F891A (PMID: 31142613).

(2) The use of deletion mutants for PCNA and RAD51 interaction is inferior to using specific point mutants, as done for the SUMO interaction and the sites for post-translational modifications.

(3) Figure 4D and Figure 5A report data with standard deviations, which is unusual for n=2. Maybe the individual data points could be plotted with a color for each independent experiment to allow the reader to evaluate the reproducibility of the results.

Comments on revisions:

In this revision, the authors adequately addressed my concerns. The only issue I see remaining is the site of Srs2 action. The authors argue in favor of gaps and against R-loops and ssDNA resulting from excessive supercoiling. The authors do not discuss ssDNA resulting from processing of onesided DSBs, which are expected to result from replication run-off after CPT damage but are not expected to provide the 3'-junction for preferred PCNA loading. Can the authors exclude PCNA at the 5'-junction at a resected DSB?

We have now added a sentence stating that we cannot exclude the possibility that PCNA may be positioned at a 5’-junction, as this can be observed in vitro, albert that PCNA loading was seen exclusively at a 3’-junction in the presence of RPA (Ellison & Stillman, 2003; Majka et al, 2006).

Recommendations For the authors:

Reviewer #2 (Recommendations For the authors):

A Bonferroni correction should be made for the multiple comparisons in several figures.

Specific comments:

l. 41. This is a too long and confusing sentence.

Sentence shortened: “These data suggest that Srs2 recruitment to PCNA proximal ssDNA-RPA filaments followed by its sumoylation can promote checkpoint recovery, whereas Srs2 action is minimized at regions with no proximal PCNA to permit RPA-mediated ssDNA protection”.

l. 60. Identify Ddc2 and Mec1 as ATRIP and ATR.

Done.

l. 125 "fails to downregulate RPA levels on chromatin and Mec1-mediated DDC..." fails to downregulate RPA and fails to reduce Mec1-mediated DDC?

Sentence modified: “fails to downregulate both the RPA levels on chromatin and the Mec1-mediated DDC”

l. 204 "consistent with the notion that Srs2 has roles beyond RPA regulation"... What other roles? It's stripping of Rad51? Removing toxic joint molecules? Something else?

Sentence modified: “consistent with the notion that Srs2 has roles beyond RPA regulation, such as in Rad51 regulation and removing DNA joint molecules”.

l. 249 "Significantly, srs2-ΔPIM and -3KR increased the percentage of rfa1-zm2 cells transitioning into the G1 phase" No. Just back to normal. As stated in l. 258: "258 We found that srs2-ΔPIM and srs2-3KR mutants on their own behaved normally in the two DDC assays described above." All of these effects are quite small.

Sentence modified: “Compared with rfa1-zm2 cells, srs2-∆PIM rfa1-zm2 and srs2-3KR rfa1-zm2 cells showed increased percentages of cells transitioning into the G1 phase”.

l. 468 "Our previous work has provided several lines of evidence to support that Rad51 removal by Srs2 is separable from the Srs2-RPA antagonism (Dhingra et al., 2021). What evidence? See my comment above about not having both proteins removed at the same time.

We have addressed this point in our initial rebuttal and some key points are summarized below. In our previous report (Dhingra et al., 2021), we provided several lines of evidence to support the conclusion that Rad51 is not relevant to the Srs2-RPA antagonism. For example, while rad51∆ rescues the hyper-recombination phenotype of srs2∆ cells, rad51∆ did not affect the hyper-checkpoint phenotype of srs2∆. In contrast, rfa1-zm1/zm2 have the opposite effects, that is, rfa1zm1/zm2 suppressed the hyper-checkpoint, but not the hyper-recombination, phenotype of srs2∆ cells. The differential effects of rad51∆ and rfa1-zm1/zm2 were also seen for the ATPase dead allele of Srs2 (srs2K41A). For example, rfa1-zm2 rescued hyper-checkpoint and CPT sensitivity of srs2-K41A cells, while rad51∆ had neither effect. These and other data described by Dhingra et al (2021) suggest that Srs2’s effects on checkpoint vs. recombination can be separated genetically. Consistent with our conclusion summarized above, deleting the Rad51 binding domain in Srs2 (srs2-∆Rad51BD) has no effect on rfa1-zm2 phenotype in CPT (Fig. 2D). This data provides yet another evidence that Srs2 regulation of Rad51 is separable from the Srs2RPA antagonism.

l. 525 "possibility, we tested the separation pin of Srs2 (Y775), which was shown to enables its in vitro helicase activity during the revision of our work..." ?? there was helicase activity during the revision of your work? Please fix the sentence.

Sentence modified: “we tested the separation pin of Srs2 (Y775). This residue was shown to be key for the Srs2’s helicase activity in vitro in a report that was published during the revision of our work (Meir et al, 2023).”

Fig. 3. "srs2-ΔPIM and -3KR allow better G1 entry of rfa1-zm2 cells." is it better entry or less arrest at G2/M? One implies better turning off of a checkpoint, the other suggests less activation of the checkpoint.

This is a correct statement. For all strains examined in Figure 3, cells were seen in G2/M phase after 1-hour CPT treatment, suggesting proper arrest.

References:

Armstrong AA, Mohideen F, Lima CD (2012) Recognition of SUMO-modified PCNA requires tandem receptor motifs in Srs2. Nature 483: 59-63

Colavito S, Macris-Kiss M, Seong C, Gleeson O, Greene EC, Klein HL, Krejci L, Sung P (2009) Functional significance of the Rad51-Srs2 complex in Rad51 presynaptic filament disruption. Nucleic Acids Res 37: 6754-6764.

De Tullio L, Kaniecki K, Kwon Y, Crickard JB, Sung P, Greene EC (2017) Yeast Srs2 helicase promotes redistribution of single-stranded DNA-bound RPA and Rad52 in homologous recombination regulation. Cell Rep 21: 570-577

Dhingra N, Kuppa S, Wei L, Pokhrel N, Baburyan S, Meng X, Antony E, Zhao X (2021) The Srs2 helicase dampens DNA damage checkpoint by recycling RPA from chromatin. Proc Natl Acad Sci U S A 118: e2020185118

Elango R, Sheng Z, Jackson J, DeCata J, Ibrahim Y, Pham NT, Liang DH, Sakofsky CJ, Vindigni A, Lobachev KS et al (2017) Break-induced replication promotes formation of lethal joint molecules dissolved by Srs2. Nat Commun 8: 1790

Ellison V, Stillman B (2003) Biochemical characterization of DNA damage checkpoint complexes: clamp loader and clamp complexes with specificity for 5' recessed DNA. PLoS Biol 1: E33

Kolesar P, Altmannova V, Silva S, Lisby M, Krejci L (2016) Pro-recombination Role of Srs2 Protein Requires SUMO (Small Ubiquitin-like Modifier) but Is Independent of PCNA (Proliferating Cell Nuclear Antigen) Interaction. J Biol Chem 291: 7594-7607.

Kolesar P, Sarangi P, Altmannova V, Zhao X, Krejci L (2012) Dual roles of the SUMO-interacting motif in the regulation of Srs2 sumoylation. Nucleic Acids Res 40: 7831-7843.

Majka J, Binz SK, Wold MS, Burgers PM (2006) Replication protein A directs loading of the DNA damage checkpoint clamp to 5'-DNA junctions. J Biol Chem 281: 27855-27861

Meir A, Raina VB, Rivera CE, Marie L, Symington LS, Greene EC (2023) The separation pin distinguishes the pro- and anti-recombinogenic functions of Saccharomyces cerevisiae Srs2. Nat Commun 14: 8144

Menin L, Ursich S, Trovesi C, Zellweger R, Lopes M, Longhese MP, Clerici M (2018) Tel1/ATM prevents degradation of replication forks that reverse after Topoisomerase poisoning. EMBO Rep 19: e45535

Minca EC, Kowalski D (2011) Replication fork stalling by bulky DNA damage: localization at active origins and checkpoint modulation. Nucleic Acids Res 39: 2610-2623

Niu H, Klein HL (2017) Multifunctional roles of Saccharomyces cerevisiae Srs2 protein in replication, recombination and repair. FEMS Yeast Res 17: fow111

Papouli E, Chen S, Davies AA, Huttner D, Krejci L, Sung P, Ulrich HD (2005) Crosstalk between SUMO and ubiquitin on PCNA is mediated by recruitment of the helicase Srs2p. Mol Cell 19: 123-133

Redon C, Pilch DR, Rogakou EP, Orr AH, Lowndes NF, Bonner WM (2003) Yeast histone 2A serine 129 is essential for the efficient repair of checkpoint-blind DNA damage. EMBO Rep 4: 678-684

Tercero JA, Longhese MP, Diffley JFX (2003) A central role for DNA replication forks in checkpoint activation and response. Mol Cell 11: 1323-1336

Vaze MB, Pellicioli A, Lee SE, Ira G, Liberi G, Arbel-Eden A, Foiani M, Haber JE (2002) Recovery from checkpointmediated arrest after repair of a double-strand break requires Srs2 helicase. Mol Cell 10: 373-385

Changes from the earlier version to the revised version

Abstract

Abstract

Earlier version: Activation of the DNA damage checkpoint upon genotoxin treatment induces a multitude of cellular changes, such as cell cycle arrest or delay, to cope with genome stress. After prolonged genotoxin treatment, the checkpoint can be downregulated to allow cell cycle and growth resumption. In yeast, downregulation of the DNA damage checkpoint requires the Srs2 DNA helicase, which removes the ssDNA binding complex RPA and the associated Mec1 checkpoint kinase from DNA, thus dampening Mec1-mediated checkpoint. However, it is unclear whether the ‘anti-checkpoint’ role of Srs2 is temporally and spatially regulated to both allow timely checkpoint termination and to prevent superfluous RPA removal. Here we address this question by examining regulatory elements of Srs2, such as its phosphorylation, sumoylation, and protein-interaction sites. Our genetic analyses and checkpoint level assessment suggest that the RPA countering role of Srs2 is promoted by Srs2 binding to PCNA, which recruits Srs2 to a subset of ssDNA regions. RPA antagonism is further fostered by Srs2 sumoylation, which we found depends on the Srs2-PCNA interaction. Srs2 sumoylation is additionally reliant on Mec1 and peaks after Mec1 activity reaches maximal levels. Based on these data, we propose a two- step model of checkpoint downregulation wherein Srs2 recruitment to PCNA proximal ssDNA- RPA filaments and subsequent sumoylation stimulated upon Mec1 hyperactivation facilitate checkpoint recovery. This model suggests that Srs2 removal of RPA is minimized at ssDNA regions with no proximal PCNA to permit RPA-mediated DNA protection at these sites.

Revised version: Activation of the DNA damage checkpoint upon genotoxin treatment induces a multitude of cellular changes to cope with genome stress. After prolonged genotoxin treatment, the checkpoint can be downregulated to allow cell cycle and growth resumption. In yeast, downregulation of the DNA damage checkpoint requires the Srs2 DNA helicase, which removes the ssDNA binding complex RPA and the associated Mec1 checkpoint kinase from DNA, thus dampening Mec1-mediated checkpoint. However, it is unclear whether the ‘anti-checkpoint’ role of Srs2 is temporally and spatially regulated to allow timely checkpoint termination while preventing superfluous RPA removal. Here we address this question by examining regulatory elements of Srs2, such as its phosphorylation, sumoylation, and protein-interaction sites. Our genetic analyses and checkpoint level assessment suggest that the RPA countering role of Srs2 is promoted by Srs2 binding to PCNA, which recruits Srs2 to a subset of ssDNA containing regions. RPA antagonism is further fostered by Srs2 sumoylation, which we found depending on the Srs2-PCNA interaction and Mec1, and peaking after Mec1 activity reaches maximal levels. These data suggest that Srs2 recruitment to PCNA proximal ssDNA-RPA filaments followed by its sumoylation can promote checkpoint recovery, whereas Srs2 action is minimized at regions with no proximal PCNA to permit RPA-mediated ssDNA protection.

Introduction

Earlier version: The DNA damage response promotes organismal survival of genotoxic stress by eliciting many cellular changes. A critical component of this response is a signaling process termed the DNA damage checkpoint (DDC) (Enoch et al, 1992; Hartwell and Weinert, 1989; Lanz et al, 2019). Upon encountering stress, apical DDC kinases are recruited to DNA lesion sites by the damage sensing proteins and are then activated by activator proteins (Usui et al, 2001; Zou and Elledge, 2003). A universal DDC sensor is the RPA complex that binds with high affinity and specificity to single-strand DNA (ssDNA), whose abundance rises under almost all types of genotoxic stress (Marechal and Zou, 2015). Using the budding yeast system as an example, RPA bound to the ssDNA filament can interact with Ddc2, the co-factor of the Mec1 checkpoint kinase, and recruit the Mec1-Ddc2 complex to DNA lesion sites (Elledge, 1996; Foiani et al, 2000). Subsequently, Mec1 is activated by activator proteins, such as the 9-1-1 complex loaded at the 5’ DNA ends at ss-dsDNA junctions. Activated Mec1 can then phosphorylate multiple substrates, among which is a key downstream effector kinase Rad53, the phosphorylation of which turns on its own kinase activity and yields further phosphorylation events (Harrison and Haber, 2006).

Revised version: The DNA damage response promotes organismal survival of genotoxic stress by eliciting many cellular changes. A critical component of this response is a signaling process termed the DNA damage checkpoint (DDC) (Enoch et al., 1992, Hartwell and Weinert, 1989, Lanz et al., 2019). Upon experiencing genotoxic stress, apical DDC kinases are recruited to DNA lesion sites by the damage sensing proteins and are then activated by activator proteins (Usui et al., 2001, Zou and Elledge, 2003). A universal DDC sensor is the RPA complex that binds with high affinity and specificity to single-strand DNA (ssDNA), whose abundance rises under almost all types of genotoxic stress (Marechal and Zou, 2015). Using the budding yeast system as an example, RPA bound to the ssDNA filament can interact with Ddc2 (homolog of the human ATRIP protein), the co-factor of the Mec1 (homolog of the human ATR protein) checkpoint kinase, and recruit the Mec1-Ddc2 complex to DNA lesion sites (Elledge, 1996, Foiani et al., 2000).

Earlier version: Collectively, the DDC kinases can modify hundreds of proteins, inducing multitude of pathway changes (Lanz et al., 2019). For example, DDC can arrest or delay the cell cycle progression to allow more time for genome repair in promoting cellular survival (Sanchez et al, 1999; Wang et al, 2001).

Revised version: Subsequently, Mec1 is activated by activator proteins, such as the 9-1-1 complex loaded at the 5’ DNA ends at ssDNA and double strand (ds) DNA junctions. Activated Mec1 can then phosphorylate multiple substrates, among which is a key downstream effector kinase Rad53, the phosphorylation of which turns on its own kinase activity and yields further phosphorylation events (Harrison and Haber, 2006). Collectively, the DDC kinases can modify hundreds of proteins, inducing multitude of pathway changes (Lanz et al., 2019). For example, DDC can arrest or delay the cell cycle progression to allow more time for genome repair (Sanchez et al., 1999, Wang et al., 2001).

While DDC activation is crucial for the cell to cope with genotoxic stress, its timely downregulation is equally vital (Waterman et al., 2020). Many changes induced by the DDC kinases need to be reversed to allow continued cell growth and recovery. For example, the DDC-mediated cell cycle arrest has to be lifted to allow proliferation and alternative DNA repair mechanisms permitted at other cell cycle phases (Pellicioli et al., 2001, Waterman et al., 2020). Thus far, DDC downregulation has been mainly studied using the model organism yeast, with several strategies being identified (Waterman et al., 2020). One of these is mediated by the DNA helicase Srs2 that targets the most upstream step of the DDC pathway through removing RPA and the associated Mec1-Ddc2 from chromatin (Dhingra et al., 2021). Importantly, the role of Srs2 in DDC dampening (anti-checkpoint) can be separated from its well-studied role in removing the recombinase Rad51 from ssDNA (anti-recombination) that modulates homologous recombination levels and outcomes (Dhingra et al., 2021). The former role of Srs2 is critical for cell survival of genotoxins, such as camptothecin (CPT) that traps Top1 on DNA (Dhingra et al., 2021). Considering that other organisms also possess DNA helicases capable of removing RPA from ssDNA (Hormeno et al., 2022, Jenkins et al., 2021, Shorrocks et al., 2021), it is plausible that an anti-checkpoint strategy involving RPA antagonism may be present beyond yeast.

Currently, little is known about how checkpoint dampening can be temporally controlled to minimize interference with the initial DDC activation that must occur in the early phase of the DNA damage response. In the case of the Srs2-mediated checkpoint regulation, it is also unclear how the process can minimize superfluous removal of RPA from ssDNA sites that unlikely induce the DDC due to lacking 5’ DNA ends at ss-dsDNA junctions for 9-1-1 loading, such as at those formed at R-loops. The main functions of RPA at these sites can be ssDNA protection and repair, rather than DDC induction, thus these sites may benefit from shielding from Srs2-mediated RPA removal. In addressing these questions, we hypothesize that Srs2-based RPA antagonism should be selective in location and timing, and that the regulatory elements within the Srs2 protein may contribute to such selectivity.

Earlier version: To examine the above hypothesis and to define Srs2 features relevant to its RPA antagonism during DDC regulation, we inspected Srs2 protein sites involved in its post-translational modifications and interactions with other proteins. Our genetic analyses of mutations affecting each of these sites suggest that while most of the examined features do not affect RPA antagonism, Srs2 binding to PCNA, which marks 3’ DNA ends of ss-dsDNA junctions that flank ssDNA gaps or tails, and Srs2 sumoylation contribute to RPA regulation in the context of DDC. Srs2 binding to PCNA is known to recruit Srs2 to replication perturbation sites, while the roles of Srs2 sumoylation were less clear (Kolesar et al, 2012; Moldovan et al, 2007; Papouli et al, 2005; Pfander et al, 2005; Saponaro et al, 2010). Unexpectedly, we found that Srs2 binding to PCNA and Srs2 sumoylation are functionally related, since PCNA binding promotes Srs2 sumoylation. Moreover, Srs2 sumoylation level peaks late in the response to genomic stress after the DDC kinases are hyper-activated, and this modification depends on the Mec1 checkpoint kinase. Collectively, our data support a model wherein Srs2 can distinguish between two types of ssDNA regions using PCNA proximity as a guide for subsequent sumoylation, RPA removal, and countering hyper-activated checkpoint.

Revised version: To examine the above hypothesis and to define Srs2 features relevant to its RPA antagonism during DDC regulation, we inspected Srs2 protein sites involved in its post-translational modifications and interactions with other proteins. Our genetic analyses of mutations affecting each of these sites suggest that while most of the examined features do not affect RPA antagonism, Srs2 binding to PCNA, which marks 3’ DNA ends of ss-dsDNA junctions that flank ssDNA gaps or tails, and Srs2 sumoylation contribute to RPA regulation in the context of DDC. Srs2 binding to PCNA is known to recruit Srs2 to replication perturbation sites, while the roles of Srs2 sumoylation were less clear (Kolesar et al., 2012, Moldovan et al., 2007, Papouli et al., 2005, Pfander et al., 2005, Saponaro et al., 2010). Unexpectedly, we found that Srs2 binding to PCNA and Srs2 sumoylation are functionally related, since PCNA binding promotes Srs2 sumoylation. Moreover, Srs2 sumoylation level peaks late in the response to genomic stress after the DDC kinases are hyper-activated, and this modification depends on the Mec1 checkpoint kinase. Collectively, our data suggest a model wherein Srs2 can distinguish between two types of ssDNA regions using PCNA proximity as a guide for subsequent sumoylation, RPA removal, and countering hyper-activated checkpoint.

Results

Srs2 regulatory features and their disabling mutants

Earlier version: To understand which Srs2 attributes are relevant to RPA regulation during the DNA damage response, we examined mutant alleles that perturb known features of the Srs2 protein. Srs2 was reported to contain a helicase domain at its N-terminal region and three protein-protein interaction sites at its C-terminal region (Fig. 1A) (Niu and Klein, 2017). We have previously shown that the Srs2 ATPase-dead allele (srs2-K41A) fails to downregulate RPA levels on chromatin and Mec1-mediated DDC, indicating a requirement of ATPase-driven stripping of RPA from ssDNA (Dhingra et al., 2021). The reported Srs2 protein-bindings sites include a Rad51 binding domain (Rad51-BD), a PCNA interaction motif (PIM), and a SUMO interaction motif (SIM) (Fig. 1A) (Colavito et al, 2009; Kolesar et al, 2016; Kolesar et al., 2012). We employed CRISPR-based editing to generate the srs2-ΔPIM (Δ1159-1163 amino acids) and the srs2-SIMmut (1170IIVID1174 to 1170AAAAD1174) alleles used in previous studies and examined an established srs2-ΔRad51BD allele (Δ875-902 amino acids) (Colavito et al., 2009; Kolesar et al., 2016; Kolesar et al., 2012). We note that despite Srs2 removal of RPA from ssDNA, an interaction between the two proteins has yet to be reported.

Revised version: To understand which Srs2 attributes are relevant to RPA regulation during the DNA damage response, we examined mutant alleles that perturb known features of the Srs2 protein. Srs2 was reported to contain a helicase domain at its N-terminal region and three protein-protein interaction sites at its C-terminal region (Fig. 1A) (Niu and Klein, 2017). We have previously shown that the Srs2 ATPase-dead allele (srs2-K41A) fails to downregulate both RPA levels on chromatin and the Mec1-mediated DDC, indicating a requirement of ATPase-driven stripping of RPA from ssDNA (Dhingra et al., 2021). The reported Srs2 protein-bindings sites include a Rad51 binding domain (Rad51-BD), a PCNA interaction motif (PIM), and a SUMO interaction motif (SIM) (Fig. 1A) (Armstrong et al., 2012, Colavito et al., 2009, Kolesar et al., 2016, Kolesar et al., 2012, Papouli et al., 2005). We employed CRISPR-based editing to generate the srs2-ΔPIM (Δ1159-1163 amino acids) and the srs2-SIMmut (1170IIVID1174 to 1170AAAAD1174) alleles used in previous studies and examined an established srs2-ΔRad51BD allele (Δ875-902 amino acids) (Armstrong et al., 2012, Colavito et al., 2009, Kolesar et al., 2016, Kolesar et al., 2012, Papouli et al., 2005). We note that despite Srs2 can remove of RPA from ssDNA, an interaction between the two proteins has yet to be reported.

Regulatory features of Srs2 and corresponding mutants.

(A) Schematic of Srs2 protein domains, regulatory features, and mutations affecting each of its features. Cdk1 phosphorylation sites depicted here include T604, S698, S879, S893, S938, S950, and S965 (Chiolo et al., 2005). Potential Mec1 phosphorylation sites include S890 and S933 (Albuquerque et al., 2015, Faca et al., 2020). Sumoylation sites mapped previously include K1081, K1089, and K1142 (Saponaro et al., 2010). Protein-interaction domains or motifs include the Rad51 binding domain (BD), PCNA Interaction Motif (PIM), and SUMO Interaction Motif (SIM) (Colavito et al., 2009, Kolesar et al., 2016, Kolesar et al., 2012). The residues for these domains are indicated. Mutant alleles disabling each of these features are included in the parentheses.

Earlier version: (B) Mutations affecting Srs2 features do not affect Srs2 protein level. Protein extracts were prepared from asynchronous cells. Srs2 was examined using anti-Srs2 antibody by immunoblotting. Srs2 levels in each mutant were first normalized to the Pgk1 loading control and then to those in wild-type cells. Mean of two biological isolates per genotype is graphed with error bars representing standard deviation (SD).

Revised version: (B) Mutations affecting Srs2 features do not affect Srs2 protein level. Protein extracts were prepared from asynchronous cells. Srs2 was examined using anti-Srs2 antibody by immunoblotting. Srs2 levels in each mutant were first normalized to the Pgk1 loading control and then to those in wild-type cells. Mean of two biological isolates per genotype is graphed with error bars representing standard deviation (SD). No statistical differences were seen when compared the Srs2 protein level between the wild-type (WT) control and each of the indicated mutants based on student t-test.

We also queried post-translational modifications reported for Srs2 by mutating the corresponding modification sites (Fig. 1A). First, seven Cdk1 phosphorylation consensus sites of Srs2 were mutated to produce the srs2-7AV (T604V, S698A, S879A, S893A, S938A, S950A, and S965A) allele (Chiolo et al., 2005), which was previously shown to affect homologous recombination. Second, three sumoylation sites of Srs2 were mutated to generate srs2-3KR (K1081R, K1089R, and K1142R) (Saponaro et al., 2010). Finally, we mutated two serine residues (S890 and S933) to alanine to generate srs2-2SA, since phosphorylation at these sites was identified in proteomic studies as dependent on both the Mec1 kinase and genotoxic treatments (Faca et al., 2020). Because the Srs2 C-terminal end contributes to the protein’s functions (Antony et al., 2009, Colavito et al., 2009), the six mutant alleles described above were not tagged with epitopes at its C-terminus. Using an anti-Srs2 specific antibody, we showed that all six alleles supported wild-type levels of Srs2 protein based on immunoblotting (Fig. 1B).

Mutating PCNA binding and sumoylation sites of Srs2 rescues rfa1-zm2 CPT sensitivity

Srs2 loss leads to persistent DDC and perturbed recombination, both of which contribute to srs2Δ sensitivity to genotoxins, including CPT and methyl methanesulfonate (MMS) (Fig. 2A) (Niu and Klein, 2017). Significantly, our recent study has shown that RPA mutants that reduce its ssDNA binding affinity rescue the persistent DDC levels in srs2Δ cells but not their recombination defects (Dhingra et al., 2021). rfa1-zm2 (N492D, K493R, K494R) is one of the RPA alleles examined that contains mutations at three DNA contacting residues in the Rfa1 subunit of RPA (Fan and Pavletich, 2012). Purified RPA complex containing these mutations exhibited a two-fold increase in dissociation constant using ssDNA as substrate, compared to the wild-type RPA complex (Dhingra et al., 2021). In cells, rfa1-zm2 supports normal Rfa1 protein level, cell growth, DNA replication (Dhingra et al., 2021). It also did not affect gene conversion, direct repeat repair or rDNA recombination, nor the hyper-recombination phenotype of srs2Δ (Dhingra et al., 2021). However, rfa1-zm2 and srs2Δ are mutually suppressive for both their genotoxic sensitivities and their DDC abnormalities (Fig. 2B, S1A) (Dhingra et al., 2021).

Srs2 PCNA binding and sumoylation are involved in suppressing rfa1 mutants

Earlier version: (A) Schematic to show that Srs2’s roles in checkpoint dampening (anti-checkpoint for simplicity) and DNA repair both contribute to genotoxin resistance.

Revised version: (A) Schematic to show that Srs2’s roles in checkpoint dampening (anti-checkpoint for simplicity) and DNA repair, both of which contribute to genotoxin resistance.

Earlier version: (B) Schematic to highlight the mutual suppression between srs2Δ and rfa1 mutants and the underlying effects on the RPA-ssDNA levels in cells.

Revised version: (B) Schematic to highlight the mutual suppressive relationship between srs2Δ and rfa1 mutants and the underlying effects on the RPA-ssDNA levels in cells.

Earlier version: (C) srs2-ΔPIM and -3KR suppress rfa1-zm2 sensitivity to CPT. A 10-fold serial dilution of cells of the indicated genotypes were spotted and growth was assessed after incubation at 30°C for 40 hours unless otherwise noted.

Revised version: (C) srs2-ΔPIM or -3KR suppresses rfa1-zm2 sensitivity to CPT. A 10-fold serial dilution of cells of the indicated genotypes were spotted and growth was assessed after incubation at 30°C for 40 hours unless otherwise noted.

Earlier version: (D) Several srs2 mutants did not suppress rfa1-zm2 sensitivity to CPT. Experiments were performed as in panel 2C.

Revised version: (D) Several srs2 mutants did not suppress rfa1-zm2 sensitivity to CPT. Experiments were performed as in panel C.

Earlier version: (E) srs2-ΔPIM and -3KR suppress rfa1-t33 sensitivity to CPT. Top: schematic of the Rfa1 protein domains and rfa1 mutant alleles. Bottom: srs2-ΔPIM and -3KR improved CPT sensitivity of rfa1-t33 cells. Experiments were performed as in panel 2C.

Revised version: (E) srs2-ΔPIM or -3KR suppresses rfa1-t33 sensitivity to CPT. Top: schematic of the Rfa1 protein domains and rfa1 mutant alleles. Bottom: srs2-ΔPIM or -3KR improved CPT sensitivity of rfa1-t33 cells. Experiments were performed as in panel C.

These data suggest that the srs2 and rfa1-zm2 genetic relationship can be used as a readout for the antagonism between Srs2 and RPA in DDC and that this functional relationship can be separable for Srs2’s anti-recombination role stemming from Rad51 regulation. Based on this rationale, we employed the genetic suppression as an initial readout in querying whether any of the six srs2 mutant alleles described above are related to RPA regulation. Given that Srs2’s role in checkpoint dampening was best examined in CPT (Dhingra et al., 2021), which induces the DNA damage checkpoint to delay G2/M progression, but does not induce the DNA replication checkpoint, experiments were mostly carried out in this condition.

Earlier version: We found that on their own, srs2-SIMmut, -ΔPIM, and -7AV showed different levels of sensitivity toward CPT at a concentration ranging from 2 to 8 µg/ml, while srs2-ΔRad51BD and - 3KR exhibited close to wild-type level of CPT resistance, which are consistent with previous reports (Colavito et al., 2009; Saponaro et al., 2010) (Fig. S1A). In addition, srs2-2SA, which has not been examined previously, did not show CPT sensitivity (Fig. S1A). Since rfa1-zm2 cells exhibited stronger CPT sensitivity than the tested srs2 alleles (Fig. S1A), we asked if the latter could suppress rfa1-zm2 CPT sensitivity. We found that both srs2-ΔPIM and -3KR improved the CPT resistance of rfa1-zm2 cells on media containing 2 µg/ml CPT (Fig. 2C).

Revised version: We found that on their own, srs2-ΔPIM, SIMmut, and -7AV showed different levels of sensitivity toward CPT at a concentration ranging from 2 to 8 µg/ml, while srs2-ΔRad51BD and -3KR exhibited close to wild-type level of CPT resistance, which are consistent with previous reports (Colavito et al., 2009, Saponaro et al., 2010) (Fig. S1A). In addition, srs2-2SA, which has not been examined previously, did not show CPT sensitivity (Fig. S1A). Since rfa1-zm2 cells exhibited stronger CPT sensitivity than the tested srs2 alleles (Fig. S1A), we asked if the latter could suppress rfa1-zm2 CPT sensitivity. We found that both srs2-ΔPIM and -3KR improved the CPT resistance of rfa1-zm2 cells on media containing 2 µg/ml CPT (Fig. 2C).

Rescue was also seen at higher CPT concentrations, with srs2-ΔPIM exhibiting a stronger effect than srs2-3KR (Fig. S1B). In contrast to srs2-ΔPIM and srs2-3KR, srs2-ΔRad51BD and -2SA showed no effect on rfa1-zm2 growth on media containing CPT, whereas srs2-7AV and -SIMmut worsened it (Fig. 2D). The differential effects seen for the srs2 alleles are consistent with the notion that Srs2 has multiple roles in cellular survival of genotoxic stress (Bronstein et al., 2018, Niu and Klein, 2017). The unique suppression seen for mutants affecting Srs2 binding to PCNA and its sumoylation suggests that these features are relevant to RPA regulation during the DNA damage response.

Additional suppressive effects conferred by srs2-ΔPIM and -3KR toward rfa1 mutants

Earlier version: We next tested whether the observed srs2-ΔPIM and -3KR suppression was specific towards rfa1-zm2 or applicable to other rfa1 mutants that are also defective in ssDNA binding. While rfa1-zm2 affects the ssDNA binding domain C (DBD-C) of Rfa1, a widely used rfa1 allele, rfa1- t33, impairs the DBD-B domain (Fig. 2E) (Umezu et al, 1998). Unlike rfa1-zm2, rfa1-t33 additionally decreases Rfa1 protein level and interferes with DNA replication and repair (Deng et al, 2014; Umezu et al., 1998). Despite having these additional defects, CPT survival of rfa1-t33 cells was also improved by srs2-ΔPIM or srs2-3KR (Fig. 2E). As seen in the case of rfa1-zm2, srs2-ΔPIM conferred better rescue than srs2-3KR toward rfa1-t33 (Fig. 2E). When comparing the two rfa1 alleles, srs2 rescue was stronger toward rfa1-zm2, likely because rfa1-t33 has other defects besides reduced ssDNA binding (Fig. 2C, 2E). In either case, srs2 mutants only conferred partial suppression, consistent with the notion that Srs2 has roles beyond RPA regulation (Bronstein et al., 2018; Niu and Klein, 2017). Nevertheless, the observed rescue of both rfa1-zm2 and rfa1-t33 provided strong genetic evidence that Srs2 binding to PCNA and sumoylation are involved the regulation of RPA during the DNA damage response.

Revised version: We next tested whether the observed srs2-ΔPIM and -3KR suppression was specific towards rfa1-zm2 or applicable to other rfa1 mutants that are also defective in ssDNA binding. While rfa1-zm2 affects the ssDNA binding domain C (DBD-C) of Rfa1, a widely used rfa1 allele, rfa1-t33, impairs the DBD-B domain (Fig. 2E) (Umezu et al., 1998). Unlike rfa1-zm2, rfa1-t33 additionally decreases Rfa1 protein level and interferes with DNA replication and repair (Deng et al., 2014, Umezu et al., 1998). Despite having these additional defects, CPT survival of rfa1-t33 cells was also improved by srs2-ΔPIM or srs2-3KR (Fig. 2E). As seen in the case of rfa1-zm2, srs2-ΔPIM conferred better rescue than srs2-3KR toward rfa1-t33 (Fig. 2E). When comparing the two rfa1 alleles, srs2 rescue was stronger toward rfa1-zm2, likely because rfa1-t33 has other defects besides reduced ssDNA binding (Fig. 2C, 2E). In either case, srs2 mutants only conferred partial suppression, consistent with the notion that Srs2 has roles beyond RPA regulation, such as in Rad51 regulation and removing DNA joint molecules (Bronstein et al., 2018, Niu and Klein, 2017). Nevertheless, the observed rescue of both rfa1-zm2 and rfa1-t33 provided strong genetic evidence that Srs2 binding to PCNA and sumoylation are involved the regulation of RPA during the DNA damage response.

Srs2 regulation of RPA during DDC is not only observed under CPT treatment, but is also seen in treatment with MMS (Dhingra et al., 2021), which induces both DDS and the DNA replication checkpoint and thus delaying both S phase progression and G2/M phase exit (Paciotti et al., 2000, Tercero et al., 2003). Consistent with observation made in CPT conditions, srs2-ΔPIM moderately improved the MMS sensitivity of rfa1-zm2 and rfa1-t33, while srs2-3KR had a small effect on the MMS sensitivity of rfa1-zm2 (Fig. S1B, S1C). The degrees of suppression seen in MMS treatment were less than those observed in CPT conditions, and we address likely reasons for this difference in Discussion. Collectively, the multiple genetic results are consistent among themselves and provide evidence that the roles of Srs2 binding to PCNA and sumoylation in RPA antagonism can be applicable to more than one genomic stress situation.

srs2-ΔPIM and -3KR rescue checkpoint abnormality of rfa1-zm2 cells

Earlier version: We next investigated whether srs2-ΔPIM and -3KR suppression of the DNA damage sensitivity of rfa1-zm2 cells pertained to DDC regulation. We employed two readouts for DDC that are widely used. First, we used the F9 antibody, which specifically recognizes phosphorylated Rad53, to measure the levels of activated Rad53 (Fiorani et al, 2008). Second, we measure cellular ability to timely exit the G2/M phase after CPT treatment using FACS analyses. We note that CPT-induced DDC leads to G2/M delay rather than an arrest as shown in previous studies (Dhingra et al., 2021; Menin et al, 2018; Redon et al, 2003). We have recently reported that srs2Δ cells exhibit increased levels of phosphorylated Rad53 and diminished ability to timely exit G2/M phase in CPT, and that both are suppressed by rfa1-zm2 (Dhingra et al., 2021; Menin et al., 2018; Redon et al., 2003). Interestingly, rfa1-zm2 leads to similar defects and these are rescued by srs2Δ (Dhingra et al., 2021) (Fig. 3A). The mutual suppression seen in these assays is consistent with that seen for CPT sensitivity (Dhingra et al., 2021). We note that several other RPA ssDNA binding mutants, such as rfa1-t33, also exhibit phenotypes indicative of increased DDC levels (Pellicioli et al., 2001; Seeber et al, 2016; Smith and Rothstein, 1995; Zou and Elledge, 2003). This is likely because compromised ssDNA protection can lead to increased levels of DNA lesions thus enhancing the activation of Mec1 and/or its homolog Tel1.

Revised version: We next investigated whether srs2-ΔPIM and -3KR suppression of the DNA damage sensitivity of rfa1-zm2 cells pertained to DDC regulation. We employed two readouts for DDC that are widely used. First, we used the F9 antibody, which specifically recognizes phosphorylated Rad53, to measure the levels of activated Rad53 (Fiorani et al., 2008). Second, we measure cellular ability to timely exit the G2/M phase after CPT treatment using FACS analyses. We note that CPT-induced DDC leads to G2/M delay rather than an arrest as shown in previous studies (Dhingra et al., 2021, Menin et al., 2018, Redon et al., 2003). We recently reported that srs2Δ cells exhibit increased levels of phosphorylated Rad53 and diminished ability to timely exit G2/M phase in CPT, and that both are suppressed by rfa1-zm2 (Dhingra et al., 2021). Interestingly, rfa1-zm2 leads to similar defects and these are rescued by srs2Δ (Dhingra et al., 2021) (Fig. 3A). The mutual suppression seen in these assays is consistent with that seen for CPT sensitivity (Dhingra et al., 2021). We note that several other RPA ssDNA binding mutants, such as rfa1-t33, also exhibit phenotypes indicative of increased DDC levels (Pellicioli et al., 2001, Seeber et al., 2016, Smith and Rothstein, 1995, Zou and Elledge, 2003). This is likely because compromised ssDNA protection can lead to increased levels of DNA lesions thus enhancing the activation of Mec1 and/or its homolog Tel1.

srs2-ΔPIM and -3KR correct checkpoint abnormalities in rfa1-zm2 cells.

(A) A table summarizing DDC levels observed in rfa1-zm2, srs2 null or partially defective mutants, and their combined mutants in this work and in a previous study (Dhingra et al., 2021) and interpretations.

Earlier version: (B) srs2-ΔPIM and -3KR reduce the levels of phosphorylated Rad53 in rfa1-zm2 cells. Protein extracts were examined after G1 cells were released into cycling in the presence of 16 ug/ml CPT for 2 hours. Phosphorylated Rad53 (active form) was detected by the F9 antibody by immunoblotting. Phosphorylated Rad53 signals in each indicated strain were compared to the Pgk1 loading control and normalized to those of rfa1-zm2. Two biological isolates per genotype were examined in two technical repeats and results are graphed with error bars representing SD. Statistical analysis was performed using pairwise Student’s t-test. *, p <0.05; ***, p <0.001; ****, p <0.0001.

Revised version: (B) srs2-ΔPIM or -3KR reduces the levels of phosphorylated Rad53 in rfa1-zm2 cells. Protein extracts were examined after G1 cells were released into cycling in the presence of 16 ug/ml CPT for 2 hours. Phosphorylated Rad53 (active form) was detected by the F9 antibody by immunoblotting. Phosphorylated Rad53 signals in each indicated strain were compared to the Pgk1 loading control and normalized to those of rfa1-zm2. Two biological isolates per genotype were examined in two technical repeats and results are graphed with error bars representing SD. Statistical analysis was performed using pairwise Student’s t-test (*, p <0.05; ***, p <0.001; ****, p <0.0001, and non-statistical significance is not marked)

Earlier version: (C) srs2-ΔPIM and -3KR allow better G1 entry of rfa1-zm2 cells. Experiments were performed as in panel A and samples were collected at indicated time points. FACS profiles of the samples are shown at the top and percentages of G1 cells after 2 hours of CPT treatment (CPT 2 h) are plotted at the bottom. Two biological isolates per genotype were examined in two technique repeats and results are graphed with error bars representing SD. Statistical analysis was performed using pairwise Student’s t-test. *, p <0.05; **, p <0.01.

Revised version: (C) srs2-ΔPIM and -3KR allow better G1 entry of rfa1-zm2 cells. Experiments were performed as in panel A and samples were collected at indicated time points. FACS profiles of the samples are shown at the top and percentages of G1 cells after 2 hours of CPT treatment (CPT 2 h) are plotted at the bottom. Two biological isolates per genotype were examined in two technique repeats and results are graphed with error bars representing SD. Statistical analysis was performed and results are depicted as in panel B.

Earlier version: We reasoned that if srs2-ΔPIM and srs2-3KR compromise Srs2’s ability to remove RPA from DNA in cells, they should alleviate rfa1-zm2 cells’ defects in ssDNA protection and the associated DDC abnormalities (Fig. 3A). Indeed, while rfa1-zm2 exhibited about 5-fold increase in Rad53 phosphorylation levels in CPT compared with wild-type cells as seen before (Dhingra et al., 2021), srs2-ΔPIM and -3KR reduced this level by about 50% and 80%, respectively (Fig. 3B). Similar to earlier data, we observed that wild-type cells arrested in G2/M after 1 hour of CPT treatment and about 25% of these cells entered G1 after another hour (Fig. 3C) (Dhingra et al., 2021). The ability to timely exit from G2/M phase was reduced in rfa1-zm2 cells so that only 13% cells moved to G1 phase after 2 h treatment with CPT, as seen before (Fig. 3C) (Dhingra et al., 2021). Significantly, srs2-ΔPIM and -3KR increased the percentage of rfa1-zm2 cells transitioning into the G1 phase (Fig. 3C). The suppressive effects seen in both assays were stronger for srs2-ΔPIM than srs2-3KR as seen in genotoxicity assays (Fig. S1B; 3B, 3C).

Revised version: We reasoned that if srs2-ΔPIM and srs2-3KR compromise Srs2’s ability to remove RPA from DNA in cells, they should alleviate rfa1-zm2 cells’ defects in ssDNA protection and the associated DDC abnormalities (Fig. 3A). Indeed, while rfa1-zm2 exhibited about 5-fold increase in Rad53 phosphorylation levels in CPT compared with wild-type cells as seen before (Dhingra et al., 2021), srs2-ΔPIM and -3KR reduced this level by about 50% and 80%, respectively (Fig. 3B). Similar to earlier data, we observed that wild-type cells arrested in G2/M after 1 hour of CPT treatment and about 25% of these cells entered G1 after another hour (Fig. 3C) (Dhingra et al., 2021). The ability to timely exit from G2/M was reduced in rfa1-zm2 cells so that only 13% cells moved to G1 after 2 h treatment with CPT, as seen before (Fig. 3C) (Dhingra et al., 2021).

Added in the revised version: Compared with rfa1-zm2 cells, srs2-ΔPIM rfa1-zm2 and srs2-3KR rfa1-zm2 cells showed increased percentages of cells transitioning into the G1 phase (Fig. 3C). The suppressive effects seen in both assays were stronger for srs2-ΔPIM than srs2-3KR as seen in genotoxicity assays (Fig. S1B; 3B, 3C). When assaying for chromatin RPA levels, we found that while srs2-Δpim showed an increase compared with the wild-type control, rfa1-zm2 reversed this effect (Fig. S2). This data supports the Srs2 and RPA antagonism. Collectively, results from the above assays are consistent with each other and suggest that Srs2 binding to PCNA and its sumoylation contribute to RPA antagonism during DDC regulation.

Removed in the revised version: Collectively, the genetic and DDC assay data are consistent with each other and suggest that Srs2 binding to PCNA and its sumoylation contribute to RPA antagonism in DDC regulation.

srs2-ΔPIM is additive with a slx4 mutant defective in DDC recovery

Earlier version: We found that srs2-ΔPIM and srs2-3KR mutants on their own behaved normally in the two DDC assays described above, which differs from srs2Δ cells that exhibit increased levels of Rad53 phosphorylation and reduced G1 entry (Fig. 3B, 3C) (Dhingra et al., 2021). One plausible interpretation of this difference is that unlike srs2Δ, these partially defective srs2 mutants only mildly impair Srs2-mediated RPA regulation during DDC downregulation, such that their defects can be compensated for by other DDC dampening systems. To test this idea, we examined the consequences of removing the Slx4-mediated anti-checkpoint function in srs2-ΔPIM cells. Slx4 is known to bind the scaffold protein Rtt107, and their complex can compete with the DDC protein Rad9 for binding to damaged chromatin, consequently disfavoring Rad9’s role in promoting Rad53 phosphorylation (Ohouo et al, 2010; Ohouo et al, 2013). Slx4 uses an Rtt107 interaction motif (RIM) to bind Rtt107 and a previously established RIM mutant, slx4RIM (T423A, T424A, and S567A), abolishes Rtt107 binding without affecting other Slx4 protein attributes (Wan et al, 2019). slx4RIM was tagged with a TAP tag at its C-terminal end and examined together with a control wherein the wild-type Slx4 was tagged in the same manner. Compared with the control, slx4RIM exhibited mild sensitivity toward CPT (Fig. 4A). Significantly, when srs2-ΔPIM was combined with slx4RIM, the double mutant exhibited stronger CPT sensitivity compared with corresponding single mutants (Fig. 4A).

Revised version: We found that srs2-ΔPIM and srs2-3KR mutants on their own behaved normally in the two DDC assays described above, which differs from srs2Δ cells that exhibit increased levels of Rad53 phosphorylation and reduced G1 entry (Fig. 3B, 3C) (Dhingra et al., 2021). One plausible interpretation of this difference is that unlike srs2Δ, these partially defective srs2 mutants only mildly impair Srs2-mediated RPA regulation during DDC downregulation, such that their defects can be compensated for by other DDC dampening systems. To test this idea, we examined the consequences of removing the Slx4-mediated anti-checkpoint function in srs2-ΔPIM cells. Slx4 is known to bind the scaffold protein Rtt107, and their complex can compete with the DDC protein Rad9 for binding to damaged chromatin, consequently disfavoring Rad9’s role in promoting Rad53 phosphorylation (Ohouo et al., 2010, Ohouo et al., 2013). Slx4 uses an Rtt107 interaction motif (RIM) to bind Rtt107 and a previously established RIM mutant, slx4RIM (T423A, T424A, and S567A), abolishes Rtt107 binding without affecting other Slx4 protein attributes (Wan et al., 2019). slx4RIMwas tagged with a TAP tag at its C-terminal end and examined together with a control wherein the wild-type Slx4 was tagged in the same manner. Compared with the control, slx4RIM exhibited mild sensitivity toward CPT (Fig. 4A). Significantly, when srs2-ΔPIM was combined with slx4RIM, the double mutant exhibited stronger CPT sensitivity compared with corresponding single mutants (Fig. 4A).

srs2-ΔPIM and slx4RIM are additively defective in DDC down-regulation.

(A) The srs2-ΔPIMΔ slx4RIM double mutant shows stronger CPT sensitivity than either single mutant. Experiments are performed as in Figure 2C.

Earlier version: (B) srs2-ΔPIM and slx4RIM are additive for increasing the level of phosphorylated Rad53 when cells are treated with CPT. Experiments were conducted and signals are quantified as those described in Figure 3B. Results from three biological isolates per genotype are graphed with error bars representing SD. Statistical analysis was performed using pairwise Student’s t- test. *, p <0.05; **, p <0.01.

Revised version: (B) srs2-ΔPIM and slx4RIM are additive for increasing the level of phosphorylated Rad53 when cells are treated with CPT. Experiments were conducted, and data are quantified as those described in Figure 3B. Results from three biological isolates per genotype are graphed with error bars representing SD. Statistical analysis was performed using pairwise Student’s t-test. *, p <0.05; **, p <0.01.

Earlier version: (C) srs2-ΔPIM and slx4RIM are additive for reducing cells exiting from G2/M arrest at 2 hours post CPT treatment. Experiments were conducted and signals are quantified as those described in Figure 3C. Results from three biological isolates per genotype are graphed with error bars representing SD. Statistical analysis was performed using pairwise Student’s t- test. *, p <0.05; **, p <0.01.

Revised version: (C) srs2-ΔPIM and slx4RIM are additive for reducing cells exiting from G2/M phase at 2 hours post CPT treatment. Experiments were conducted and data are quantified as those described in Figure 3C. Results from three biological isolates per genotype are graphed with error bars representing SD. Statistical analysis was performed using pairwise Student’s t-test. *, p <0.05; **, p <0.01.

Earlier version: We further examined DDC levels when cells were exposed to CPT using both Rad53 phosphorylation and ability to timely exit G2/M as readouts. Compared with Slx4-TAP cells, slx4RIM-TAP cells exhibited 2-fold increase in Rad53 phosphorylation level, consistent with a role of Slx4 in downregulating DDC (Fig. 4B). In Slx4-TAP background, srs2-ΔPIM led to 3-fold increase in Rad53 phosphorylation level (Fig. 4B). As cells containing untagged Slx4 and srs2-ΔPIM did not show Rad53 phosphorylation changes (Fig. 3B), tagging Slx4 likely sensitizes srs2-ΔPIM. Significantly, the double mutant of srs2-ΔPIM and slx4RIM-TAP exhibited further increase in Rad53 phosphorylation (∼5.5 fold) compared with srs2-ΔPIM and slx4RIM-TAP single mutant (Fig. 4B). Congruous with this finding, the srs2-ΔPIM slx4RIM-TAP double mutant also reduced the percentage of cells entering G1 compared with their single mutants after 2-hour treatment with CPT (Fig. 4C). Together, these three lines of evidence support the notion that defects in the Srs2-mediated DDC downregulation can be compensated for by Slx4-mediated mechanisms.

Revised version: We further examined DDC levels when cells were exposed to CPT using both Rad53 phosphorylation and ability to timely exit G2/M as readouts. Compared with SLX4-TAP cells, slx4RIM-TAP cells exhibited a 2-fold increase in Rad53 phosphorylation levels, consistent with a role of Slx4 in downregulating DDC (Fig. 4B). In SLX4-TAP background, srs2-ΔPIM led to 2-fold increase in Rad53 phosphorylation level (Fig. 4B). As cells containing untagged Slx4 and srs2-ΔPIM did not show Rad53 phosphorylation changes (Fig. 3B), tagging Slx4 appeared to sensitize srs2-ΔPIM. Significantly, the srs2-ΔPIM slx4RIM-TAP double mutant exhibited further increase in Rad53 phosphorylation than either srs2-ΔPIM or slx4RIM-TAP single mutant (Fig. 4B). Congruous with this finding, the srs2-ΔPIM slx4RIM-TAP double mutant also reduced the percentage of cells entering G1 compared with their single mutants after 2-hour treatment with CPT (Fig. 4C). Together, these three lines of evidence support the notion that defects in the Srs2-mediated DDC downregulation can be compensated for by Slx4-mediated mechanisms.

Srs2 binding to PCNA facilitates Srs2 sumoylation

Earlier version: We next addressed whether the two Srs2 features involved in RPA antagonism are functionally related. To this end, we tested if Srs2 binding to PCNA affected Srs2 sumoylation in cells. We confirmed that endogenous sumoylation of Srs2 could be detected in cell extracts after enriching SUMOylated proteins using 8His-tagged SUMO (8His-Smt3) (Fig. 5A). Briefly, proteins were extracted from cells under denaturing conditions to disrupt protein-protein association and prevent desumoylation, and sumoylated proteins were enriched on Ni-NTA (nickel-nitrilotriacetic acid) resin (Ulrich and Davies, 2009). Consistent with previous findings, sumoylated forms of Srs2 in the Ni-NTA eluate were detected on immunoblots using an anti- Srs2 antibody only when cells contained 8His-Smt3, where unmodified Srs2 was seen in the presence or absence of 8His-Smt3 due to nonspecific binding to the resin (Fig. S2A) (Saponaro et al., 2010). These bands were only seen in wild-type cells but not srs2Δ cells (Fig. S2A). We found that Srs2 sumoylation level was reduced in srs2-3KR strains in CPT conditions, as expected from mutating the three major sumoylation sites (Fig. 5A) (Saponaro et al., 2010).

Revised version: We next addressed whether the two Srs2 features involved in RPA antagonism are functionally related. To this end, we tested if Srs2 binding to PCNA affected Srs2 sumoylation in cells. We confirmed that endogenous sumoylation of Srs2 could be detected in cell extracts after enriching SUMOylated proteins using 8His-tagged SUMO (8His-Smt3) (Fig. 5A). Briefly, proteins were extracted from cells under denaturing conditions to disrupt protein-protein association and prevent desumoylation, and sumoylated proteins were enriched on Ni-NTA (nickel-nitrilotriacetic acid) resin (Ulrich and Davies, 2009). Consistent with previous findings, sumoylated forms of Srs2 in the Ni-NTA eluate were detected on immunoblots using an anti-Srs2 antibody only when cells contained 8His-Smt3, where unmodified Srs2 was seen in the presence or absence of 8His-Smt3 due to nonspecific binding to the resin (Fig. S3A) (Saponaro et al., 2010). These bands were only seen in wild-type cells but not srs2Δ cells (Fig. S3A). We found that Srs2 sumoylation level was reduced in srs2-3KR strains in CPT conditions, as expected from mutating the three major sumoylation sites (Fig. 5A) (Saponaro et al., 2010). Interestingly, srs2-ΔPIM cells exhibited a similar reduction in Srs2 sumoylation as seen in srs2-3KR cells (Fig. 5A). This result suggests that Srs2 sumoylation largely depends on Srs2 binding to PCNA, supporting the notion that the two events occur sequentially.

Srs2 sumoylation depends on Srs2 PIM motif and the Mec1 kinase

Earlier version: (A) Srs2 sumoylation level is reduced in srs2-ΔPIM cells. Samples were collected after CPT treatment as in Figure 3A and sumoylated proteins were enriched as described in the text and methods. Sumoylated form of Srs2 (Srs2-Su) migrated slower than unmodified Srs2 on gel and both were detected using anti-Srs2 antibody in immunoblotting. Equal loading is indicated by Ponceau-S stain (Stain). The levels of sumoylated Srs2 relative to those of unmodified Srs2 were normalized to wild-type cells and plotted. Two biological isolates per genotype were examined in two technique repeats and results are graphed with error bars representing SD. Statistical analysis was performed using pairwise Student’s t-test. ns, not significant; *, p <0.05.

Revised version: (A) Srs2 sumoylation level is reduced in srs2-ΔPIM cells. Samples were collected after CPT treatment as in Figure 3A and sumoylated proteins were enriched as described in the text and Methods. Sumoylated form of Srs2 (Srs2-Su) migrated slower than unmodified Srs2 on gel and both were detected using anti-Srs2 antibody in immunoblotting. Equal loading is indicated by Ponceau-S stain (Stain). The levels of sumoylated Srs2 relative to those of unmodified Srs2 were normalized to wild-type cells and plotted. Two biological isolates per genotype were examined in two technique repeats and results are graphed with error bars representing SD. Statistical analysis was performed using pairwise Student’s t-test. *, p<0.05.

(B) Srs2 sumoylation level during CPT treatment. Samples were collected as described in Figure 3B. Experiments were conducted, and immunoblotting data is presented as those in panel A. Time 0 samples are G1 cells before CPT treatment. FACS analysis at corresponding time points are included below the plot.

(C) Rad9 phosphorylation level during CPT treatment. Samples used in panel B were examined for Rad9 phosphorylation using an anti-Rad9 antibody. Unmodified and phosphorylated (Rad9-p) Rad9 forms are indicated. FACS analysis at corresponding time points are indicated below blots.

(D) Quantification of Srs2 sumoylation and Rad9 phosphorylation level changes during CPT treatment. Srs2 sumoylation signals were normalized based on unmodified form. Values at each timepoint were relative to that of G1 sample (Time 0) before CPT treatment, and fold changes were indicated on left y-axis. Rad9 signals were normalized to non-modified form and percentage of phosphorylated form were indicated on right y-axis. Results from three biological isolates per genotype are graphed.

(E) Mec1 is required for Srs2 sumoylation. Experiments were performed and data are presented as in panel A.

Removed in the revised version: Interestingly, srs2-ΔPIM cells exhibited a similar reduction in Srs2 sumoylation as seen in srs2- 3KR cells (Fig. 5A). This result suggests that Srs2 sumoylation largely depends on Srs2 binding to PCNA, supporting the notion that the two events occur sequentially.

Srs2 sumoylation peaks after maximal Mec1 activation and depends on Mec1

Earlier version: While RPA can engage with ssDNA generated in different processes, such as transcription, DNA repair, and DNA replication, only a subset of ssDNA regions is adjacent to 3’-junctions loaded with PCNA. This is because PCNA loading requires an ssDNA and dsDNA junction site with a 3’ DNA end and occurs mainly during S phase (Moldovan et al., 2007). We thus reasoned that Srs2 binding to PCNA may prevent Srs2 from removing RPA bound to ssDNA that forms transiently or is not flanked by PCNA, such as within R-loop regions and negatively supercoiling DNA regions, both of which are generated abundantly in CPT conditions due to topological stress (Pommier et al, 2022). This could provide a means for spatial control of Srs2’s anti-RPA role. We explored whether this role could also be regulated temporally. As the genetic data described above suggest that Srs2 sumoylation acts downstream of Srs2 binding to PCNA in antagonizing RPA, we examined the timing of Srs2 sumoylation after CPT treatment.

Revised version: While RPA can engage with ssDNA generated in different processes, such as transcription, DNA repair, and DNA replication, only a subset of ssDNA regions is adjacent to 3’-junctions loaded with PCNA. This is because PCNA loading requires an ssDNA and dsDNA junction site with a 3’ DNA end (Moldovan et al., 2007). We thus reasoned that Srs2 binding to PCNA may prevent Srs2 from removing RPA bound to ssDNA that is not flanked by PCNA, such as within R-loop regions and negatively supercoiling DNA regions, both of which are generated abundantly in CPT conditions due to topological stress (Pommier et al., 2022). This could provide a means for spatial control of Srs2’s anti-RPA role. We explored whether this role could also be regulated temporally. As the genetic data described above suggest that Srs2 sumoylation acts downstream of Srs2 binding to PCNA in antagonizing RPA, we examined the timing of Srs2 sumoylation after CPT treatment.

Removed in the revised version: We released G1 synchronized cells into media containing CPT and examined Srs2 sumoylation at several time points (Fig. 5B). A similar experimental scheme was used to monitor the level of Mec1 activation using Rad9 phosphorylation as a readout (Fig. 5C).

Earlier version: Quantification of Srs2 sumoylated forms relative to its unmodified forms showed that its sumoylation level peaked between 1.5 and 2 h after cells were released into CPT (Fig. 5D). Quantification of phosphorylated forms of Rad9 relative to its unmodified form showed that checkpoint activation peaked 1 hour after cell release (Fig. 5D). That checkpoint level peaked before Srs2 sumoylation level peaked suggests a possible dependence of the latter on the checkpoint. In testing this idea, we found that mec1Δ cells, which contain sml1Δ to support viability (Zhao et al, 1998), abolished Srs2 sumoylation (Fig. 5E). Neither sml1Δ alone nor removal of Rad53 affected Srs2 sumoylation (Fig. 5E), suggesting that Mec1’s effect on this modification does not require the downstream Rad53 effector kinase. As Mec1 is not generally required for protein sumoylation during the DNA damage response (Cremona et al, 2012), its involvement in Srs2 sumoylation presents a rather unique effect.

Revised version: We released G1 synchronized cells into media containing CPT and examined Srs2 sumoylation at several time points (Fig. 5B). A similar experimental scheme was used to monitor the level of Mec1 activation using Rad9 phosphorylation as a readout (Fig. 5C). Quantification of Srs2 sumoylated forms relative to its unmodified forms showed that its sumoylation level peaked between 1.5 and 2 h after cells were released into CPT (Fig. 5D). Quantification of phosphorylated forms of Rad9 relative to its unmodified form showed that checkpoint activation peaked 1 hour after cell release (Fig. 5D). That checkpoint level peaked before Srs2 sumoylation level peaked suggests a possible dependence of the latter on the checkpoint. In testing this idea, we found that mec1Δ cells, which contain sml1Δ to support viability (Zhao et al., 1998), abolished Srs2 sumoylation (Fig. 5E). Neither sml1Δ alone nor removal of Rad53 affected Srs2 sumoylation (Fig. 5E), suggesting that Mec1’s effect on this modification does not require the downstream Rad53 effector kinase. As Mec1 is not generally required for protein sumoylation during the DNA damage response (Cremona et al., 2012), its involvement in Srs2 sumoylation presents a rather unique effect.

Examination of Mec1-S1964 phosphorylation in Srs2-mediated DDC regulation

The combined results that Srs2 sumoylation levels peak after Rad9 phosphorylation reaches maximal level and that Srs2 sumoylation depends on Mec1 raised the possibility that Mec1 hyper-activation during the late part of the DNA damage response may favor Srs2’s role in DDC dampening. Recently, Mec1 protein bound to Ddc2 was shown to be auto-phosphorylated at Ser1964 in a late part of the DNA damage response after a single DNA break is induced by the HO endonuclease (Memisoglu et al., 2019). Further, the phosphorylation-defective mutant mec1-S1964A impaired checkpoint downregulation in this condition, leading to the model that Mec1-S1964 phosphorylation contributes to DDC dampening at least after HO-induced break (Memisoglu et al., 2019). We thus tested whether Mec1-S1964 phosphorylation is also involved in Srs2-mediated DDC regulation under CPT conditions.

Earlier version: First, we asked whether Mec1-S1964 phosphorylation occurs under CPT conditions using a time course scheme similar to the one used in testing Srs2 sumoylation as described above. Ddc2-bound Mec1 was recovered after immunoprecipitating Ddc2 at each timepoint and the proteins were examined by immunoblotting. To detect Mec1-S1964 phosphorylation, we used an antibody raised against the peptide containing this modification provided kindly by the Haber lab (Memisoglu et al., 2019). The antibody detected the Mec1 S1964 phosphorylation band in immunoprecipitated fractions only when Ddc2 was tagged with the Myc tag but not when Ddc2 was untagged (Fig. S2B). Further, the Mec1 S1964 phosphorylation band was largely abolished in the mec1-S1964A cells. (Fig. S2B; 6A). Given that the antibody also detects a closely spaced background band (Fig. S2B; 6A), we calculated the relative level of Mec1 S1964 phosphorylation in each sample by normalizing to this band. Quantification results showed that Mec1 S1964 phosphorylation level peaked around the same time as seen for Rad9 phosphorylation, which was 1 hour after CPT treatment (Fig. 6A).

Revised version: First, we asked whether Mec1-S1964 phosphorylation occurs under CPT conditions using a time course scheme similar to the one used in testing Srs2 sumoylation as described above. Ddc2-bound Mec1 was recovered after immunoprecipitating Ddc2 at each timepoint and the proteins were examined by immunoblotting. To detect Mec1-S1964 phosphorylation, we used an antibody raised against the peptide containing this modification provided kindly by the Haber lab (Memisoglu et al., 2019). The antibody detected the Mec1 S1964 phosphorylation band in immunoprecipitated fractions only when Ddc2 was Myc tagged but not when Ddc2 was untagged (Fig. S3B). Further, the Mec1 S1964 phosphorylation band was largely abolished in the mec1-S1964A cells (Fig. S3B; 6A). Given that the antibody also detects a closely spaced background band (Fig. S3B; 6A), we calculated the relative level of Mec1 S1964 phosphorylation in each sample by normalizing to this band. Quantification results showed that Mec1 S1964 phosphorylation level peaked around the same time as seen for Rad9 phosphorylation, which was 1 hour after CPT treatment (Fig. 6A).

Mec1 and RPA phosphorylation did not link to Srs2 anti-checkpoint function.

(A) Mec1 autophosphorylation at S1964 during CPT treatment. Samples were collected as described in Figure 3B. Ddc2-myc was immunoprecipitated and the associated Mec1 was co-purified. The anti-Mec1-S1964-p antibody detected the phosphorylated form of Mec1 and a non-specific band (*). Quantification below the blot was based on results from two biological isolates.

(B) mec1-S1964A did not rescue rfa1-zm2 sensitivity toward CPT. Experiments were performed as in Figure 2C.

(C) Examination of the rfa1-S178A and -S178D mutants. Top: schematic of the Rfa1 protein domains and position of Ser178. ssDNA binding domain (DBD) A-C are shown. Bottom, rfa1-S178A and -S178D had no effect on srs2Δ’s sensitivity to CPT. Experiments were performed as in Figure 2C.

(D) Examination of the rfa2-3SA and -3SD mutants. Top, schematic of the Rfa2 protein domains and position of Ser122, Ser187 and Ser189 mutated in rfa2-3SA and -3SD. ssDNA binding domain (DBD)-D is indicated. Bottom, rfa2-3SA has no effect on srs2Δ’s sensitivity toward CPT while rfa2-3SD showed additive effect. Experiments were performed as in Figure 2C.

Earlier version: After confirming that Mec1 S1964 phosphorylation also took place under CPT conditions, we tested whether the mec1-S1964A or the phosphorylation-mimetic mutant, mec1- S1964E, influenced rfa1-zm2 or srs2 growth on media containing CPT. If Mec1 S1964 phosphorylation aids Srs2-mediated RPA antagonism in CPT, we would expect mec1-S1964A to behave similarly to srs2 mutants in rescuing rfa1-zm2 CPT sensitivity, while mec1-S1964E should have the opposite effects. We found that mec1-S1964A, which showed normal CPT resistance on its own, did not affect the growth of rfa1-zm2 cells on CPT-containing media (Fig. 6B). mec1-S1964A also did not influence the growth of srs2-ΔPIM and srs2-3KR cells regardless of rfa1-zm2 (Fig. 6B; S3A).

Revised version: After confirming that Mec1 S1964 phosphorylation also took place under CPT conditions, we tested whether mec1-S1964A or the phosphorylation-mimetic mutant, mec1-S1964E, influenced rfa1-zm2 or srs2 growth on media containing CPT. If Mec1 S1964 phosphorylation aids Srs2-mediated RPA antagonism in CPT, we would expect mec1-S1964A to behave similarly to srs2 mutants in rescuing rfa1-zm2 CPT sensitivity, while mec1-S1964E should have the opposite effects. We found that mec1-S1964A, which showed normal CPT resistance on its own, did not affect the growth of rfa1-zm2 cells on CPT-containing media (Fig. 6B). mec1-S1964A also did not influence the growth of srs2-ΔPIM and srs2-3KR cells regardless of rfa1-zm2 (Fig. 6B; S4A).

Earlier version: Though mec1-S1964E on its own did not cause CPT sensitivity, it led to slight but reproducible slow growth in cells containing srs2-3KR and srs2-ΔPIM with or without rfa1-zm2 (Fig. S3B). We found that neither mec1-S1964A nor mec1-S1964E affected Srs2 sumoylation levels (Fig. S3C). Thus, Mec1, but not its S1964 phosphorylation, is required for Srs2 sumoylation. These data provide evidence that even though an HO-induced DNA break and CPT treatment both induce Mec1 S1964 phosphorylation, this modification appears to not affect Srs2-mediated DDC dampening in the latter situation.

Revised version: Though mec1-S1964E on its own did not cause CPT sensitivity, it led to slight but reproducible slow growth in cells containing srs2-3KR and srs2-ΔPIM with or without rfa1-zm2 (Fig. S4B). We found that neither mec1-S1964A nor mec1-S1964E affected Srs2 sumoylation levels (Fig. S4C). Thus, Mec1, but not its S1964 phosphorylation, is required for Srs2 sumoylation. These data provide evidence that even though an HO-induced DNA break and CPT treatment both induce Mec1 S1964 phosphorylation, this modification appears to not affect Srs2-mediated DDC dampening in the latter situation.

Genetic data suggest that known RPA phosphorylation sites do not affect DDC recovery

Earlier version: While we mainly focused on the identification of Srs2 features that contribute to RPA antagonism during DDC regulation, we also considered the possibility that certain regulatory RPA features may contribute to this control as well. In particular, RPA is known to be phosphorylated by Mec1 (Chen et al, 2010; Faca et al., 2020; Lanz et al, 2021). We thus tested whether Mec1-mediated phosphorylation may render RPA more susceptible to be removed from ssDNA by Srs2. To this end, we mutated previously identified Mec1 phosphorylation sites on the Rfa1 and Rfa2 subunits. These include Rfa1 Ser178 and three sites on Rfa2 (S122, S187, S189) (Albuquerque et al, 2015; Chen et al., 2010). We generated corresponding non- phosphorylatable mutants, rfa1-S178A and rfa2-3SA (S122A, S187A, and S189A), and phosphomimetic mutants, rfa1-S178D and rfa2-3SD.

Revised version: While we mainly focused on the identification of Srs2 features that contribute to RPA antagonism during DDC regulation, we also considered the possibility that certain regulatory RPA features may contribute to this control as well. RPA is known to be phosphorylated by Mec1 (Chen et al., 2010, Faca et al., 2020, Lanz et al., 2021). We thus tested whether Mec1-mediated phosphorylation may render RPA more susceptible to be removed from ssDNA by Srs2. To this end, we mutated previously identified Mec1 phosphorylation sites on the Rfa1 and Rfa2 subunits. These include Rfa1 Ser178 and three sites on Rfa2 (S122, S187, S189) (Albuquerque et al., 2015, Chen et al., 2010). We generated corresponding non-phosphorylatable mutants, rfa1-S178A and rfa2-3SA (S122A, S187A, and S189A), and phosphomimetic mutants, rfa1-S178D and rfa2-3SD.

Using genetic suppression as a readout for possible involvement of RPA mutants in Srs2-based DDC regulation, we examined the interactions between the mutants affecting RPA phosphorylation sites as described above and srs2Δ. If RPA phosphorylation promote its removal by Srs2, the SD mutants may show suppressive interactions with srs2Δ, while SA mutants would behave like srs2 mutants. We found that neither rfa1-S178D nor S178A, both of which showed normal CPT resistance on their own, affected srs2Δ cell growth on CPT-containing media (Fig. 6C). In contrast, rfa2-3SD or rfa2-3SA, which also behaved like wild-type regarding CPT sensitivity, sensitized srs2Δ growth on CPT-containing media (Fig. 6D). This negative genetic interaction contrasts with the positive interaction seen between srs2Δ and rfa1-zm2 or rfa1-t33 (Fig. S1A) (Dhingra et al., 2021). The differential Srs2 genetic interactions among distinct RPA alleles likely reflect the multifunctional natures of Srs2 and RPA. Collectively, these analyses suggest that the known phosphorylation sites on RPA are unlikely to be involved in Srs2-based DDC regulation.

Discussion

Earlier version: Timely termination of DDC is crucial for cell survival, yet the underlying mechanisms are not well understood. Using the budding yeast model system, we investigated how the Srs2-mediated DDC dampening through antagonizing RPA can be regulated. Our previous work reports a requirement of Srs2’s ATPase activity in this role; here we examined additional Srs2 features involved in either binding to other proteins or post-translational modifications. Our results suggest that among the examined features, Srs2 binding to PCNA and its sumoylation are involved in RPA antagonism during DDC downregulation. We showed that mutants impairing these features, namely srs2Δ-PIM and srs2-3KR, suppressed both DDC abnormalities and CPT sensitivity of rfa1-zm2 cells. Interestingly, these features are functionally linked, since optimal Srs2 sumoylation requires its binding to PCNA. Further, the Srs2 sumoylation level peaks after Mec1 activity reaches its maximum and that Mec1 is required for this modification. We posit a model that can interpret previous and current data. In this model, Srs2 removal of RPA is favored at a subset of ssDNA regions that have proximal PCNA (Fig. 7, top). PCNA recruits Srs2 toward the adjacent RPA-ssDNA filament, and promotes Srs2 sumoylation together with hyperactive Mec1, and both effects foster downregulation of the Mec1 checkpoint (Fig. 7, bottom).

Revised version: Timely termination of DDC is crucial for cell survival, yet the underlying mechanisms are not well understood. Here we address this question in the context of Srs2-mediated DDC dampening in budding yeast. Our previous work reports a requirement of Srs2’s ATPase activity in this role; here we examined additional Srs2 features involved in either binding to other proteins or post-translational modifications. Our results suggest that among the examined features, Srs2 binding to PCNA and its sumoylation are involved in RPA antagonism during DDC downregulation. We showed that mutants impairing these features, namely srs2Δ-PIM and srs2-3KR, suppressed both DDC abnormalities and CPT sensitivity of rfa1-zm2 cells. Interestingly, these features are functionally linked, since optimal Srs2 sumoylation requires its binding to PCNA. Further, the Srs2 sumoylation level peaks after Mec1 activity reaches its maximum and that Mec1 is required for this modification. We posit a model that can interpret previous and current data. In this model, Srs2 removal of RPA is favored at a subset of ssDNA-containing regions that have proximal PCNA (Fig. 7, top). PCNA recruits Srs2 toward the adjacent RPA-ssDNA filament, and promotes Srs2 sumoylation together with hyperactive Mec1 to foster downregulation of the Mec1 checkpoint (Fig. 7, bottom).

A working model for the regulation of Srs2-mediated checkpoint dampening.

RPA-coated ssDNA regions generated from multiple DNA metabolic processes include those without proximal PCNA, such as ssDNA within R-loops and negatively supercoiled regions, and those with proximal PCNA, such as ssDNA gaps associated with perturbed DNA replication or repair. In latter situations, the resultant RPA-ssDNA can recruit the Mec1/Ddc2 kinase complex. Mec1 is subsequently activated by other DDC proteins, such as the 9-1-1 complex that demarcates the 5’ end junctions adjacent to ssDNA regions (not drawn). After DDC is activated for a prolonged period, Srs2 recruited to PCNA that demarcates the 3’ end junction flanking the ssDNA regions can be sumoylated. Srs2 sumoylation level increases as Mec1 activation heightens and can result in more efficient RPA removal, leading to Mec1 downregulation.

Added in the revised version: RPA-ssDNA filaments can be separated into two groups from the perspective of PCNA loading status. The first group contains 3’ DNA end at the ss-dsDNA junction that permit PCNA loading (Fig. 7, top, green). These regions can include ssDNA gaps produced from perturbed replication (Sogo et al., 2002) or from excision and repair of trapped Top1 (Sun et al., 2020).

Added in the revised version: Due to PCNA proximity, this group of ssDNA regions can favor Srs2 recruitment and RPA removal. Since they often also contain 5’ DNA ends required for loading the 9-1-1 complex (Majka et al., 2006), these sites can be ideal for Srs2’s action in checkpoint dampening.

Removed in the revised version: RPA-ssDNA filaments can be separated into two groups based on PCNA loading status

Earlier version: The first group contains 3’ DNA end at the ss-dsDNA junction to permit PCNA loading (Fig. 7, top, green). These regions can include ssDNA gaps produced from perturbed replication (Sogo et al, 2002)or from excision and repair of trapped Top1 (Sun et al, 2020). Due to PCNA proximity, this group of ssDNA regions can favor Srs2 recruitment and RPA removal. Since they often also contain 5’ DNA ends required for loading the 9-1-1 complex (Majka et al, 2006), these sites can be ideal for Srs2’s action in checkpoint dampening. The second group of ssDNA regions are devoid of 3’ DNA end thus PCNA loading. These regions can be within R-loops or negatively supercoiled regions, which are induced by CPT treatment due to increased topological stress associated with the depletion of the functional pool of Top1 (Pommier et al., 2022). These regions are unlikely to contain 5’ DNA ends at ss-dsDNA junctions; thus RPA there mainly supports DNA protection and repair and not DDC. We reason that these ssDNA sites can be shielded from Srs2 removal of RPA to allow better repair and protection. The notion that ssDNA regions close to PCNA are preferred for Srs2 anti-checkpoint action provides a rationale for how Srs2 could remove RPA at some sites and minimize unnecessary RPA loss from other sites. Quantitative assessment of RPA residence on DNA and genome-wide mapping of Srs2 action sites in the future can further test the proposed models and provide more clarity. Regardless, results presented in this work support the conclusion that Srs2 can distinguish between two types of ssDNA regions using PCNA proximity as a guide for RPA removal.

Revised version: Biochemical data suggested that PCNA is loaded exclusively at the 3’ DNA ends at ssDNA-dsDNA junctions by the RFC complex in the presence of RPA, though PCNA loading could occur at 5’ DNA ends in the absence of RPA (Ellison and Stillman, 2003, Majka et al., 2006). We thus do not exclude the possibility that ssDNA-dsDNA junctions containing 5’ DNA ends, such as those formed at resected DNA breaks, may also allow Srs2-based RPA regulation. The second group of ssDNA regions are devoid of DNA ends thus PCNA loading (Fig. 7, top, white). These regions can be within R-loops or negatively supercoiled regions, which are induced by CPT treatment due to increased topological stress associated with the depletion of the functional pool of Top1 (Pommier et al., 2022). These regions are unlikely to contain DNA ends at ss-dsDNA junctions; thus, RPA there mainly supports DNA protection and repair and not DDC. We reason that these ssDNA sites can be shielded from Srs2 removal of RPA to allow better repair and protection. The notion that ssDNA regions close to PCNA are preferred for Srs2 anti-checkpoint action provides a rationale for how Srs2 could remove RPA at some sites while minimizing unnecessary RPA loss from other sites. Quantitative assessment of RPA residence on DNA and genome-wide mapping of Srs2 action sites in the future can further test the proposed models and provide more clarity. Regardless, results presented in this work provide evidence for the model that Srs2 can distinguish between two types of ssDNA regions using PCNA proximity as a guide for RPA removal.

Earlier version: Genetic analysis showed that srs2Δ-PIM also suppressed rfa1-zm2 sensitivity toward MMS, though to a less degree as seen in CPT conditions. While more than one factor could contribute to this difference, we note that CPT only activates the DNA damage checkpoint, while MMS additionally induces DNA replication checkpoint (Menin et al., 2018; Redon et al., 2003; Tercero et al., 2003). It is thus possible that the Srs2-RPA antagonism is more important for the DNA damage checkpoint compared with the DNA replication checkpoint. Further investigation of this possibility among others will shed light on differential suppressive effects seen here.

Revised version: Genetic analysis showed that srs2Δ-PIM also suppressed rfa1-zm2 sensitivity toward MMS, though to a less degree as seen in CPT conditions. While more than one factor could contribute to this difference, we note that CPT only activates the DNA damage checkpoint, while MMS additionally induces the DNA replication checkpoint (Menin et al., 2018, Redon et al., 2003, Tercero et al., 2003). It is thus possible that the Srs2-RPA antagonism is more important for the DNA damage checkpoint compared with the DNA replication checkpoint. With this said, we do not exclude the possibility that the Srs2-RPA antagonism may also be relevant to specific forms of DNA repair. Further investigation of this possibility among others will shed light on differential suppressive effects seen here.

Earlier version: An involvement of PCNA binding in Srs2-based RPA regulation could compensate for a lack of Srs2 interaction with RPA. This scenario differs from Srs2-based Rad51 stripping that requires their interaction (Antony et al., 2009; Colavito et al., 2009; Krejci et al, 2003). It is thus likely that though both the anti-recombinase and anti-checkpoint roles of Srs2 benefit from selectivity, they are achieved via different means. Our previous work has provided several lines of evidence to support that Rad51 removal by Srs2 is separable from the Srs2-RPA antagonism (Dhingra et al., 2021). Consistent with this conclusion, deleting the Rad51 binding domain in Srs2 had no effect on rfa1-zm2 phenotype in CPT (Fig. 2D). This is in sharp contrast to mutating the PCNA binding and the sumoylation sites of Srs2, which suppressed rfa1-zm2 defects (Fig. 2C). This data provides additional evidence that Srs2 regulation of Rad51 is separable from the Srs2-RPA antagonism. For this reason, the Srs2-based Rad51 function was not further examined in this study. However, future studies to address how Srs2 regulates RPA and Rad51 in different manners and whether there is a crosstalk between them in specific contexts will help to generate a more comprehensive understanding of the Srs2’s anti-checkpoint vs. anti- recombination functions.

Revised version: An involvement of Srs2-PCNA binding in Srs2-based RPA regulation could compensate for a lack of interaction between Srs2 and RPA. This scenario differs from Srs2-based Rad51 stripping that requires Srs2 interacts with Rad51 (Antony et al., 2009, Colavito et al., 2009, Krejci et al., 2003). It is thus likely that though both the anti-recombinase and anti-checkpoint roles of Srs2 benefit from selectivity of protein to be stripped from DNA, they are achieved via different means. Our previous work has provided several lines of evidence to support that Rad51 removal by Srs2 can be functionally separable from the Srs2-RPA antagonism (Dhingra et al., 2021). Consistent with this conclusion, deleting the Rad51 binding domain in Srs2 had no effect on rfa1-zm2 phenotype in CPT (Fig. 2D). This contrasts with mutating the PCNA binding and the sumoylation sites of Srs2, which suppressed rfa1-zm2 defects (Fig. 2C). This data provides additional evidence that Srs2 regulation of Rad51 is separable from the Srs2-RPA antagonism at genetic level. For this reason, the Srs2-based Rad51 function was not further examined in this study. However, future studies to address how Srs2 regulates RPA and Rad51 in different manners and whether there is a crosstalk between them in specific contexts will help to generate a more comprehensive understanding of both roles of Srs2.

Earlier version: Previous studies have suggested that both PIM and SIM aid Srs2 recruitment to sumoylated form of PCNA (Papouli et al., 2005; Pfander et al., 2005). Our data reveal that at least in CPT conditions, Srs2-PIM, but not Srs2-SIM, is involved in RPA regulation. This observation is congruent with phenotypic differences between Srs2’s PIM and SIM seen in other studies and here, such as much stronger genotoxic sensitivity of the latter (Fig. S1A) (Fan et al, 2023; Kolesar et al., 2016; Kolesar et al., 2012). Considering that only a small percentage of PCNA is sumoylated and mainly during S phase (Papouli et al., 2005; Pfander et al., 2005), it is probably advantageous for Srs2 not being restricted to sumoylated PCNA in order to achieve efficient RPA regulation.

Revised version: Previous studies have suggested that both PIM and SIM aid Srs2 recruitment to sumoylated form of PCNA (Papouli et al., 2005, Pfander et al., 2005). Our data reveal that at least in CPT conditions, Srs2-PIM, but not Srs2-SIM, is involved in RPA regulation. This observation is congruent with phenotypic differences between Srs2’s PIM and SIM seen in other studies and here, such as a much stronger genotoxic sensitivity of the latter (Fig. S1A) (Fan et al., 2023, Kolesar et al., 2016, Kolesar et al., 2012). Considering that only a small percentage of PCNA is sumoylated and mainly during S phase (Papouli et al., 2005, Pfander et al., 2005), it is probably advantageous for Srs2 not being restricted to sumoylated PCNA in order to achieve efficiency in RPA regulation.

Earlier version: How could Srs2 binding to PCNA helps Srs2 sumoylation? While a full answer awaits future studies, we speculate that this interaction may render Srs2 more permissive for sumoylation and/or position Srs2 in proximity with its sumoylation E3 ligase Siz2 that binds RPA (Chung and Zhao, 2015; Kolesar et al., 2012). Considering that Mec1, Ddc2, and RPA contain SUMO interaction motifs (Psakhye and Jentsch, 2012), it is possible that Srs2 sumoylation could help enrich the helicase at ssDNA-RPA sites via interacting with these motifs, though other effects are also possible. Compared with srs2-3KR, srs2-ΔPIM showed better suppression of rfa1-zm2 defects. This indicates that during RPA regulation, PCNA plays additional roles besides promoting Srs2 sumoylation.

Revised version: How could Srs2 binding to PCNA helps Srs2 sumoylation? While a full answer awaits future studies, we speculate that this interaction may render Srs2 more permissive for sumoylation and/or position Srs2 in proximity with its sumoylation E3 ligase Siz2 that binds RPA (Chung and Zhao, 2015, Kolesar et al., 2012). Considering that Mec1, Ddc2, and RPA contain SUMO interaction motifs (SIMs) (Psakhye and Jentsch, 2012), it is possible that sumoylation could help enrich Srs2 at ssDNA-RPA sites via SIM interactions, though other effects are possible. Compared with srs2-3KR, srs2-ΔPIM showed better suppression of rfa1-zm2 defects. This indicates that during RPA regulation, PCNA plays additional roles besides promoting Srs2 sumoylation.

Earlier version: We found that Srs2 sumoylation in cells also requires Mec1, which does not promote sumoylation in general (Cremona et al., 2012). On the contrary, Mec1 loss causes increased sumoylation levels (Cremona et al., 2012). It is currently unclear how Mec1 can specifically promote Srs2 sumoylation, but this may involve Mec1-mediated phosphorylation of SUMO substrates and/or enzymes. A requirement of Mec1 for Srs2 sumoylation and the dynamic nature of sumoylation suggest that Srs2 sumoylation may be a potential timing regulator. It is possible that hyperactivation of Mec1 favors Srs2 sumoylation that may aid its own downregulation. While fully testing this idea requires time-course phosphor-proteomic studies under genotoxin treatment, we examined a Mec1 auto-phosphorylation site implicated in its downregulation after induction of a single DNA break (Memisoglu et al., 2019). We did not find evidence that it is involved in Mec1 dampening in CPT condition, raising the possibility that distinct mechanisms may be used in different genomic stress conditions. We also found that two phosphorylation sites on Srs2 that showed Mec1 dependency in proteomic studies had no effect in the CPT situation (Albuquerque et al., 2015; Faca et al., 2020); however we note that these sites do not fit with the typical S/TQ Mec1 consensus sites, thus may not be the bona fide Mec1 sites. Other phosphorylation sites examined here include five Mec1 sites on Rfa1 and Rfa2 (Albuquerque et al., 2015; Chen et al., 2010). Genetic analyses did not support their roles in the Srs2 and RPA antagonism. Thus, future studies to comprehensively map DDC kinase phosphorylation sites on Srs2 and RPA, as well as their regulators, such as the RPA chaperone Rtt105 (Li et al, 2018), will be needed to test how the kinases can communicate with Srs2- based RPA regulation during the DNA damage response.

Revised version: We found that Srs2 sumoylation in cells also requires Mec1, which does not promote sumoylation in general (Cremona et al., 2012). On the contrary, Mec1 loss causes increased sumoylation levels of multiple proteins (Cremona et al., 2012). It is currently unclear how Mec1 can specifically promote Srs2 sumoylation, but this may involve Mec1-mediated phosphorylation of SUMO substrates and/or enzymes. A requirement of Mec1 for Srs2 sumoylation and the dynamic nature of sumoylation suggest that Srs2 sumoylation may be a potential timing regulator. It is possible that hyperactivation of Mec1 favors Srs2 sumoylation that may aid its own downregulation. While fully testing this idea requires time-course phosphor-proteomic studies under genotoxin treatment, we examined a Mec1 auto-phosphorylation site implicated in its downregulation after induction of a single DNA break (Memisoglu et al., 2019). We did not find evidence that it is involved in Mec1 dampening in CPT condition, raising the possibility that distinct mechanisms may be used in different genomic stress conditions. We also found that two phosphorylation sites on Srs2 that showed Mec1 dependency in proteomic studies had no effect in the CPT situation (Albuquerque et al., 2015, Faca et al., 2020); however we note that these sites do not fit with the typical S/TQ Mec1 consensus sites, thus may not be the bona fide Mec1 sites. Other phosphorylation sites examined here include five Mec1 sites on Rfa1 and Rfa2 (Albuquerque et al., 2015, Chen et al., 2010). Genetic analyses did not support their roles in the Srs2 and RPA antagonism. Thus, comprehensively map DDC kinase phosphorylation sites on Srs2, RPA, and possibly their regulators, such as the RPA chaperone Rtt105 (Li et al., 2018), will be needed to identify targets of the kinases that can facilitate Srs2-based RPA regulation.

Added in the revised version: While we provide multiple lines of evidence in support of the involvement of Srs2 binding to PCNA and its sumoylation in RPA regulation during DDC downregulation, srs2-ΔPIM or -3KR alone behaved normally in checkpoint assays. A lack of defect here can be due to buffering effects, both from not-yet characterized Srs2 features and from other checkpoint dampening pathways. In addressing the former possibility, we tested the separation pin of Srs2 (Y775).

Earlier version: While we provide multiple lines of evidence in support of the involvement of Srs2 binding to PCNA and its sumoylation in RPA regulation during DDC downregulation, srs2-ΔPIM or -3KR alone behaved normally in checkpoint assays. A lack of defect here can be due to buffering effects, both from not-yet characterized Srs2 features and from other checkpoint dampening pathways. In addressing the former possibility, we tested the separation pin of Srs2 (Y775), which was shown to enables its in vitro helicase activity during the revision of our work (Meir et al, 2023). We found that the srs2-Y775A mutation that severely compromised the Srs2 helicase activity rescued rfa2-zm2 sensitivity to CPT (Fig. S4). This result suggests that Y775 may directly contribute to in vivo RPA removal; future examination of this feature, alone and in conjunction with PCNA binding and sumoylation, can extend our understanding of Srs2- mediated DDC regulation. We also provided evidence for the latter possibility, since srs2-ΔPIM or -3KR showed additive phenotype when combined with a slx4 mutant defective in DDC dampening function. The presence of multiple checkpoint dampening factors that target different checkpoint modules highlights the importance of the process. Recent demonstration of other DNA helicases, such as the human HELB, HELQ, and BLM proteins, in stripping of RPA from ssDNA in vitro suggests that Srs2-like checkpoint downregulation may present in human cells as well (Hormeno et al., 2022; Jenkins et al., 2021; Shorrocks et al., 2021). Since RPA-ssDNA binding not only plays a key role in checkpoint, but also in many other processes, the Srs2 features uncovered here may help to better understand how RPA dynamics is controlled by helicases and other factors during additional processes beyond DDC.

Revised version: This residue was shown to be key for the Srs2’s helicase activity in vitro in a report that was published during the revision of our work (Meir et al., 2023). We found that srs2-Y775A that severely compromised the Srs2 helicase activity rescued rfa2-zm2 sensitivity to CPT (Fig. S5). This result suggests that Y775 could contribute to in vivo RPA removal; future examination of this feature in conjunction with PCNA binding and sumoylation, can extend our understanding of Srs2-mediated DDC regulation. We also provided evidence for the latter possibility, since srs2-ΔPIM or -3KR showed additive phenotype when combined with a slx4 mutant defective in DDC dampening function. The presence of multiple checkpoint-dampening factors that target different checkpoint modules highlights the importance of the process. Recent demonstration of other DNA helicases, such as the human HELB, HELQ, and BLM proteins, in stripping of RPA from ssDNA in vitro suggests that Srs2-like checkpoint downregulation may present in human cells as well (Hormeno et al., 2022, Jenkins et al., 2021, Shorrocks et al., 2021). Since RPA-ssDNA binding not only plays a key role in checkpoint, but also in many other processes, the Srs2 features uncovered here may help to better understand how RPA dynamics is controlled by helicases and other factors during additional processes beyond DDC.

Methods

Yeast strains and genetic techniques

Earlier version: Standard procedures were used for cell growth and media preparation. The strains used in this work are listed in Table S1and are isogenic to W1588-4C, a RAD5 derivative of W303 (Zhao and Blobel, 2005). Mutant alleles of srs2, rfa1, rfa2 and mec1 used in this work were generated by CRISPR-Cas9-based gene replacement, except for srs2-3KR that was introduced using a PCR-based method (DiCarlo et al, 2013). All mutations were verified by sequencing. Standard procedures were used for tetrad analyses and genotoxin testing. For each assay, at least two biological duplicates were used per each genotype. Cells were grown at 30 °C unless otherwise stated.

Revised version: Standard procedures were used for cell growth and media preparation. The strains used in this work are listed in Table S1 and are isogenic to W1588-4C, a RAD5 derivative of W303 (Zhao and Blobel, 2005). Mutant alleles of srs2, rfa1, rfa2 and mec1 used in this work were generated by CRISPR-Cas9-based gene replacement, except for srs2-3KR that was introduced using a PCR-based method (DiCarlo et al., 2013). All mutations were verified by sequencing. Standard procedures were used for tetrad analyses and genotoxin testing. For each assay, at least two biological duplicates were used per each genotype. Cells were grown at 30 °C unless otherwise stated.

Cell synchronization and cell cycle analyses

Cells from log-phase cultures were treated with 5 μg/mL α factor (GenScript RP01002) for 1 h, followed by an additional 2.5 μg/mL α factor for 30 min. When at least 95% cells were arrested in G1 phase based on the percentage of unbudded cells, they were released into yeast extract–peptone–dextrose (YPD) media containing 100 μg/mL Protease (Millipore 53702) and 16 μg/mL CPT (Sigma C9911) for 2 h. Cell cycle progression was monitored by standard flow cytometry analyses as described previously (Zhao and Rothstein, 2002). To derive percentage of G1 cells in Figure 3B, the G1 region of the “G1 sample” was used to demarcate the G1 region of the “CPT 2h” sample of the same strain in the same experiment.

Protein extraction using trichloroacetic acid (TCA)

To examine the protein levels of Srs2, the phosphorylation form of Rad53, and Rad9, cell extracts were prepared as reported (Regan-Mochrie et al., 2022) . Briefly, 2 × 108 cells were collected at indicated time points. Cell pellets were resuspended in 20% TCA and lysed by glass bead beating. The lysate was then centrifuged to remove supernatant. Precipitated proteins were suspended in Laemmli buffer (65 mM Tris-HCl at pH 6.8, 2% SDS, 5% 2-mercaptoethanol, 10% glycerol, 0.025% bromophenol blue) with 2 M Tris added to neutralize the solution. Prior to loading, samples were boiled for 5 min and spun down at 14,000 rpm for 10 min to remove insoluble materials. Samples were separated on 4-20% Tris-glycine gels (Bio-Rad 456-1096) to examine Srs2 level and phosphorated Rad53 form or 7% Tris-acetate gels (Thermo Fisher EA03555) to detect Rad9 phosphorylation.

Detection of Srs2 sumoylation

Sumoylated proteins were pulled down from cells containing 8His-tagged SUMO expressed from its endogenous locus using the standard Ni-NTA pull-down method as previously described (Ulrich and Davies, 2009). Briefly, protein extracts prepared in 55% TCA were incubated in buffer A (6 M guanidine HCl, 100 mM sodium phosphate at pH 8.0, 10 mM Tris-HCl at pH 8.0) with rotation for 1 h at room temperature. The cleared supernatant was obtained after centrifuging for 20 min and was then incubated overnight at room temperature with Ni-NTA resin (Qiagen 30210) in the presence of 0.05% Tween-20 and 4.4 mM imidazole with rotation. Beads were washed twice with buffer A supplemented with 0.05% Tween 20 and then four times with buffer C (8 M urea, 100 nM sodium phosphate at pH 6.3, 10 mM Tris-HCl at pH 6.3) supplemented with 0.05% Tween 20. Proteins were eluted from the beads using HU buffer (8 M urea, 200 mM Tris-HCl at pH 6.8, 1 mM EDTA, 5% SDS, 0.1% bromophenol blue, 1.5% DTT, 200 mM imidazole). Samples were loaded onto NuPAGETM 3– 8% Tris-acetate gels (Thermo Fisher EA03752) for immunoblotting to detect both sumoylated and unmodified Srs2. Equal loading was verified using Ponceau S staining.

Detection of Mec1 S1964 phosphorylation during CPT treatment

G1-arrested cells were filtered and resuspended in prewarmed YPD media containing 16 μg/mL CPT to allow entry into S phase. Samples were collected at indicated time points. Cells were disrupted by glass bead beating in lysis buffer, followed by centrifugation at 20,000 g for 15 min to obtain whole-cell extract. The lysis buffer consisted of 25 mM K-HEPES (pH 7.6), 100 mM NaCl, 100 mM K-glutamate, 5 mM Mg(OAc)2, 0.02% NP40, and 0.5% Triton X-100 supplemented by protease inhibitor cocktail (Sigma P8215) and Complete Ultra EDTA-free protease inhibitor (Roche 11836145001). The activated form of Mec1 was enriched by immunoprecipitating myc-tagged Ddc2 using Protein A beads and an anti-Myc antibody (Thermo Fisher 9E10) for 2–4 h at 4°C. Beads were washed five to six times with lysis buffer, and proteins were eluted using Laemmli buffer (65 mM Tris-HCl at pH 6.8, 2% SDS, 5% 2-mercaptoethanol, 10% glycerol, 0.025% bromophenol blue). After boiling for 5 min, eluted proteins were loaded onto NuPAGETM 3–8% Tris-acetate gels for SDS-PAGE and subsequent immunoblotting analyses.

Immunoblotting analysis and antibodies

Earlier version: After SDS-PAGE, proteins were transferred to a 0.2-μm nitrocellulose membrane (GE, #G5678144) for immunoblotting. Pgk1 was used as a loading control. Antibodies used were anti-myc (Thermo, Fisher 9E10). anti-Srs2 (Santa Cruz, yC-18), anti-Pgk1 (Invitrogen, 22C5D8), F9 (a kind gift from Marco Foiani and Daniele Piccini, The FIRC Institute of Molecular Oncology, Milan, Italy) (Pellicioli et al., 2001), anti-Rad9 (a kind gift from John Petrini, MSKCC, NY, USA) (Usui et al, 2009), anti-Mec1-S1964-p (a kind gift from James E. Haber, Brandeis University, MA, USA) (Memisoglu et al., 2019). Validation of antibodies is provided either by the manufacturers or in the cited references. Membranes were scanned with a Fujifilm LAS-3000 luminescent image analyzer, which has a linear dynamic range of 104. The signal intensities of non-saturated bands were measured using ImageJ software. For graphs, data are shown as mean and SD. Statistical differences were determined using two-sided Student’s t tests.

Revised version: After SDS-PAGE, proteins were transferred to a 0.2-μm nitrocellulose membrane (GE, #G5678144) for immunoblotting. Pgk1 was used as a loading control. Antibodies used were anti-myc (Thermo, Fisher 9E10). anti-Srs2 (Santa Cruz, yC-18), anti-Pgk1 (Invitrogen, 22C5D8), F9 (a kind gift from Marco Foiani and Daniele Piccini, The FIRC Institute of Molecular Oncology, Milan, Italy) (Pellicioli et al., 2001), anti-Rad9 (a kind gift from John Petrini, MSKCC, NY, USA) (Usui et al., 2009), anti-Mec1-S1964-p (a kind gift from James E. Haber, Brandeis University, MA, USA) (Memisoglu et al., 2019). Validation of antibodies is provided either by the manufacturers or in the cited references. Membranes were scanned with a Fujifilm LAS-3000 luminescent image analyzer, which has a linear dynamic range of 104. The signal intensities of non-saturated bands were measured using ImageJ software. For graphs, mean and SD are depicted.

Added in the revised version: Statistical analyses

Added in the revised version: Pairwise comparison was conducted, and statistical differences were determined using two-sided Student’s t-test as indicted in the graphs and figure legends.